TLR4 signaling drives tissue inflammation, Claudin-5 internalization, and vascular barrier breakdown in a mouse model of neonatal meningitis

Author response:

The following is the authors’ response to the previous reviews

We are grateful to you and the reviewers for your careful and positive assessment. In response to the comment below from reviewer #2, we have expanded Figure 1- figure supplement 1 to show the results (images plus quantification) of ERG and PU.1 immunostaining together with DAPI staining that underly the calculation of the percent on non-endothelial cells that are recombined by Cdh5-CreER.

However, I found their explanation as to how they arrived at only 18% of the fibroblasts being recombined confusing. Mostly because it looks like there are many GFP+/ERG- cells in panels A & D, these would be the recombined fibroblasts and AB cells. Considering that Cdh5 gene is pretty broadly expressed across LPM, the prediction is the recombination rate would be higher (recognizing mice strains vary and this is inducible Cre).

Ideally, they would perform recombination analysis with a TF expressed by all fibroblasts in combination with Erg (like Foxc1). However, a potentially simpler approach could be to quantify this using DAPI and Erg in existing images, of the total DAPI+, what are GFP+/DAPI+/Erg- (fibroblasts) vs GFP+/ERG+/DAPI+ (endothelial). The figure would be improved by adding DAPI to one set of panels with GFP/ERG, this would show a lot of DAPI+/GFP- cells, encompassing CD206+ cells and non-recombined fibroblasts.

Thank you for overseeing this manuscript.

Changes from the earlier version to the revised version

Abstract

Neonatal bacterial meningitis is a leading cause of infant morbidity and mortality, yet the molecular and cellular basis of the leptomeningeal response to infection remains poorly defined. Here, we study a mouse model of neonatal E. coli meningitis, combining conditional gene knockouts, leptomeningeal single-nucleus RNA sequencing, and endothelial cell culture to explore the role of Toll-like receptor 4 (TLR4) signaling in the host response to infection. Deletion of Tlr4 in non-myeloid cells dramatically reduced the inflammatory response in all leptomeningeal cell types and abrogated the infection- associated increase in vascular permeability. In a brain endothelial cell line (bEnd.3 cells), exposure to E. coli triggered NF-κB activation, selective internalization of Claudin- 5, and increased monolayer permeability, responses that were eliminated by Tlr4 knockout. RNA-seq showed that TLR4 controls an NF-κB–driven transcriptional program that orchestrates the endothelial response to E. coli. These findings reveal multiple TLR4-dependent host responses to neonatal Gram-negative bacterial meningitis.

Introduction

Bacterial meningitis is a major global health challenge, causing ∼250,000 deaths annually (Brouwer, 2010; GBD 2016 Neurology Collaborators, 2019). Among neonates, Group B Streptococcus (Streptococcus agalactiae, GBS) and Escherichia coli K1 are the predominant causes of meningitis and are responsible for roughly 40% and 30% of cases, respectively (Ouchenir et al., 2017; Bundy et al., 2023). While vaccination has markedly reduced mortality in childhood meningitis, neonatal infections continue to carry high case-fatality rates, especially in low-income regions, and roughly half of the survivors experience lasting neurological deficits (Edmond et al., 2010; Guarnera et al., 2024). The high susceptibility of neonates likely reflects the immaturity of both the adaptive immune system and multiple CNS barriers, including the blood–brain barrier (BBB), the blood–cerebrospinal fluid (CSF) barrier, and the leptomeninges (Saunders et al., 2012; Yang et al., 2023). Understanding host–pathogen interactions in this developmental window is essential for designing strategies to protect the neonatal brain (Kim, 2003).

The leptomeninges form a surface barrier enveloping the brain and spinal cord. It contains a dense vascular network and large numbers of immune cells within the CSF-filled subarachnoid space. Within the leptomeninges, endothelial cells (ECs) possess BBB properties, including expression of Claudin-5 (Cldn5), Occludin, and transporters characteristic of CNS microvessels (Mastorakos et al., 2019). At present, the response of leptomeningeal ECs to bacterial infection, including alterations in barrier properties, is largely unexplored.

The present study focuses on the role of one innate immune receptor, Toll-like receptor 4 (TLR4), in regulating the leptomeningeal response to bacterial meningitis. TLR4 is a pattern- recognition receptor that detects lipopolysaccharide (LPS), the major outer membrane component of Gram-negative bacteria such as E. coli (Medzhitov, 2001; Kawai and Akira, 2007). Activation of TLR4 initiates MyD88- and TRIF-dependent signaling cascades, leading to the induction of proinflammatory cytokines, chemokines, and adhesion molecules that orchestrate leukocyte recruitment to infected tissues (Vallières and Rivest, 1997; Pålsson-McDermott and O’Neill, 2004; Nagyoszi et al. 2010; Casanova et al., 2011). In the context of bacterial meningitis, TLR4 has been shown to mediate both host-protective and pathological responses, including vascular inflammation and leukocyte infiltration (Van Well et al., 2013; Doran et al., 2016; Too et al., 2019; Wang et al., 2023). While these studies established TLR4 as a key driver of meningitis-associated inflammation, the molecular and cellular events associated with the TLR4-dependent host response during bacterial meningitis remain largely unexplored.

Claudin-5 (Cldn5) is a prominent tight-junction protein in brain ECs. Originally identified as an endothelial-specific component of tight junctions (Morita et al., 1999), Cldn5 was later shown to be indispensable for restricting paracellular permeability in vivo, as Cldn5-deficient mice exhibit size-selective leakage of small molecules across the BBB (Nitta et al., 2003). Under physiological conditions, Cldn5 is localized to endothelial cell–cell contacts, but its distribution and stability are dynamically regulated by endocytosis, recycling, and degradation pathways that respond to inflammatory stimuli (Engelhardt and Sorokin, 2009; Coisne and Engelhardt, 2011; Hashimoto et al., 2023). During CNS inflammation and infection, pro- inflammatory cytokines and innate immune receptor activation disrupt tight-junction continuity and alter endothelial barrier properties. This junctional remodeling includes Cldn5 internalization, but the upstream signaling pathways controlling Cldn5 trafficking remain incompletely understood.

A recent single-cell transcriptome study defined the genomic response of the leptomeninges in a mouse model of neonatal E. coli meningitis (Wang et al., 2023). Leptomeningeal ECs exhibited a clear inflammatory transcriptional signature, implicating the leptomeningeal vasculature as a locus of immune signaling during E. coli infection. Whether these signals directly control tight junction remodeling and vascular leak was not tested.

Here, we study a mouse neonatal meningitis model and cultured CNS ECs (bEnd.3 cells) to test whether non-myeloid TLR4 signaling orchestrates the multi-cellular response of the leptomeninges to infection and regulates Cldn5 trafficking and vascular barrier function. In the present work, we (i) used single nucleus (sn)RNA-seq to profile cell-type–specific responses in the leptomeninges from WT mice, endothelial and fibroblast (i.e., non-myeloid) Tlr4 conditional knockout mice (Tlr4VEKO), and myeloid Tlr4 conditional knockout mice (Tlr4MKO); (ii) quantified intravascular tracer leakage and immune pathway activation in vivo; (iii) tracked Cldn5 redistribution and NF-κB nuclear translocation in bEnd.3 cells exposed to E. coli; (iv) generated CRISPR Tlr4KO bEnd.3 lines to assess the TLR4 dependence of NF-κB activation, tight-junction dynamics, and barrier permeability; (v) quantified Cldn5 colocalization with junctional, plasma- membrane, and endosomal/lysosomal markers; and (vi) defined TLR4-dependent transcriptional programs in bEnd.3 cells exposed to E. coli. The results of these experiments establish a framework for understanding the pathophysiology of neonatal meningitis that establishes TLR4 signaling as a central regulator of both the multi-cellular leptomeningeal response to infection and Cldn5 trafficking and junctional integrity.

Results

Cell-type-specific transcriptional responses of the leptomeninges from WT, Tlr4VEKO, and Tlr4MKO mice to neonatal E. coli infection

Moving from brain to skull, the meninges consist of four distinct layers: (1) the pia, a thin, semi-permeable layer that mediates solute exchange between the CNS and CSF; (2) the subarachnoid space, a highly vascularized region that contains a high density of immune cells; (3) the arachnoid and its barrier layer, demarcating the boundary between CNS and non-CNS territories; and (4) the dura, a fibrous layer outside of the CNS with a high density of blood vessels, lymphatics, and immune cells (Figure 1A). The first three layers constitute the leptomeninges. Taking advantage of a natural cleavage plane between dura and leptomeninges, we have optimized whole mount immunostaining and confocal imaging of the leptomeninges together with the adjacent cortex, preserving the anatomical integrity of the pia and subarachnoid space (Figure 1B).

Cell-type-specific transcriptional responses of leptomeninges from WT, Tlr4VEKO, and Tlr4MKO mice to neonatal E. coli infection.

(A) Schematic of the leptomeninges (pia, subarachnoid space, and arachnoid barrier) and adjacent structures. (B) Immunostaining of P6 mouse leptomeninges shown as a whole mount (left) or as a transverse section through a series of Z-planes (right). The vertical black bars indicate the extent of the leptomeninges. Scale bars, 20 µm. (C) UMAP plot of leptomeninges snRNA-seq from WT, Tlr4VEKO, and Tlr4MKO mice at P6, from uninfected and E. coli-infected mice, with cell clusters identified. N=126,235 nuclei. Fb, fibroblast. (D) UMAP comparison of P6 leptomeningeal snRNA-seq from control vs. 24-hour E. coli- infected mice of the indicated genotypes. (E) snRNA-seq from four leptomeninges cell clusters (Fibroblast-arachnoid, Dural border cells, Fibroblast-pia, and Arachnoid barrier cells) showing differentially expressed genes in control vs. E. coli-infected mice of the indicated genotypes at P6. For each cell cluster, DEGs were chosen based on a comparison of infected vs. uninfected WT cells. (F and G) As in panel (E) except for ECs (F) and myeloid cells (G). DEGs, differentially expressed genes. (H) As in panel (F) and (G) except DEGs were related to NF-κB and TNF-α signaling. (I) Summary of transcriptional responses of leptomeninges from WT, Tlr4VEKO, and Tlr4MKO mice to E. coli infection.

In the E. coli meningitis model used here, postnatal day (P)5 mice were given a subcutaneous injection of 20 µl of phosphate buffered saline (PBS) with 120,000 colony forming units of E. coli RS218 (O18:K1:H7), a clinical isolate from the CSF of a neonate with bacterial meningitis (Zhu et al., 2020; Wang et al., 2023) or, as a control, 20 µl of phosphate buffered saline (PBS). This E. coli strain was additionally transformed with a plasmid encoding a red fluorescent protein (RFP) to facilitate detection of bacteria in infected tissues. Infected and control mice were sacrificed at P6, 24 hours after inoculation. P5 mice were chosen to approximate the developmental stage of the human newborn BBB and immune system, as described in Wang et al. (2023). This model recapitulates the main temporal sequence of events in human neonates with bacterial meningitis: peripheral entry of bacteria (e.g., via skin abrasions during birth), hematogenous spread of bacteria (bacteremia), and bacterial seeding of the meninges.

To assess cell type-specific roles of TLR4 signaling, we used two conditional knockout lines: Tlr4VEKO (=Cdh5-CreER;Tlr4fl/−) and Tlr4MKO (=Lyz2Cre/+;Tlr4fl/−). Cdh5 encodes Vascular Endothelial (VE) Cadherin, and the Cdh5-CreER line carries a Cdh5 promotor-CreER transgene (Monvoisin et al., 2006). Lyz2, also referred to as LysM, encodes a myeloid-specific lysozyme, and the Lyz2Cre line is a Cre knock-in at the Lyz2 locus (Clausen et al., 1999). The cell type specificities of Cdh5-CreER and Lyz2Cre were assessed using a Rosa26 loxP-stop-loxP reporter (Rosa26-LSL-tdT-2A-nlsGFP; Wang et al., 2018). The specificity of Cdh5-CreER was additionally assessed using a Rosa26-LSL-SUN1-sfGFP reporter (Mo et al., 2015).

Earlier version: Throughout most of the CNS, the Cdh5-CreER line directs Cre-mediated recombination specifically to ECs (e.g., Wang et al., 2025b). However, in the meninges, only ∼75% of the cells recombined by the Cdh5-CreER transgene are ECs and ∼25% are neither ECs nor myeloid cells (Figure 1 – figure supplement 1). Quantification of DAPI stained sections of leptomeninges shows that ECs (ERG+ cells) constitute 31+/-9% of leptomeningeal cells, and myeloid cells (PU.1+ cells) constitute 14+/-5% of leptomeningeal cells. Therefore, dural border cells, arachnoid barrier cells, and fibroblasts together constitute the remaining ∼55% of leptomeningeal cells. Since the non-ECs and non-myeloid cells that express Cdh5-CreER are one-third as numerous as ECs, they constitute ∼10% of leptomeningeal cells, equivalent to ∼18% of the combined population of dural border cells, arachnoid barrier cells, and fibroblasts. These data imply that at P2, when 4HT was administered, the Cdh5-CreER transgene produced Cre-mediated recombination in all or nearly all ECs, in a minority of dural border cells, arachnoid barrier cells, and fibroblasts, and not in myeloid cells, consistent with the pattern of Cdh5 expression in leptomeninges described by others (Mapunda et al., 2023; Pietilä et al., 2023; Smyth et al., 2024). Extrapolating these data to Tlr4, Tlr4VEKO mice should be considered to have complete or nearly complete Tlr4 knockout in ECs, with Tlr4 additionally knocked out in a fraction of the other non-myeloid cell types. For simplicity, we will refer to the Tlr4VEKO KO pattern as “non-myeloid” KO of Tlr4, although the data in Figure 1 – figure supplement 1 imply that it is predominantly an EC KO.

Revised version: Throughout most of the CNS, the Cdh5-CreER line directs Cre-mediated recombination specifically to ECs (e.g., Wang et al., 2025b). However, in the meninges, only ∼75% of the cells recombined by the Cdh5-CreER transgene are ECs and ∼25% are neither ECs nor myeloid cells (Figure 1 – figure supplement 1A, C, and D). Quantification of DAPI stained sections of leptomeninges shows that ECs (ERG+ cells) constitute 31+/-9% of leptomeningeal cells, and myeloid cells (PU.1+ cells) constitute 14+/-5% of leptomeningeal cells (Figure 1 – figure supplement 1B, C, and E). Therefore, dural border cells, arachnoid barrier cells, and fibroblasts together constitute the remaining ∼55% of leptomeningeal cells. Since the non-ECs and non-myeloid cells that express Cdh5-CreER are one-third as numerous as ECs, they constitute ∼10% of leptomeningeal cells, equivalent to ∼18% of the combined population of dural border cells, arachnoid barrier cells, and fibroblasts. These data imply that at P2, when the single dose of 4HT was administered, the Cdh5-CreER transgene produced Cre-mediated recombination in all or nearly all ECs, in a minority of dural border cells, arachnoid barrier cells, and fibroblasts, and not in myeloid cells, consistent with the pattern of Cdh5 expression in leptomeninges described by others (Mapunda et al., 2023; Pietilä et al., 2023; Smyth et al., 2024). Extrapolating these data to Tlr4, Tlr4VEKO mice should be considered to have complete or nearly complete Tlr4 knockout in ECs, with Tlr4 additionally knocked out in a fraction of the other non-myeloid cell types. For simplicity, we will refer to the Tlr4VEKOKO pattern as “non- myeloid” KO of Tlr4, although the data in Figure 1 – figure supplement 1 imply that it is predominantly an EC KO.

Earlier version: Expression of the Lyz2Cre knock-in allele was restricted to myeloid cells within the leptomeninges as determined by immunostaining for three partially overlapping myeloid markers, CD206, ASC, and PU.1 (Figure 1 – figure supplement 2). Lyz2Cre directed Cre-recombination with 90-100% efficiency in CD206+ cells and with 50-70% efficiency in ASC+ and PU.1+ cells, the range depending on whether tdTomato or GFP colocalization was being scored (Figure 1 – figure supplement 2). In the text that follows, we will refer to Lyz2Cre-recombined cells simply as “myeloid cells”, although they should be understood as CD206+ myeloid cells. The low abundance of TLR4 and the limitations of commercial anti-TLR4 antibodies precluded a direct immunohistochemical assessment of TLR4 loss in Tlr4VEKO and Tlr4MKO mice.

Revised version: Expression of the Lyz2Cre knock-in allele was restricted to myeloid cells within the leptomeninges as determined by immunostaining for three partially overlapping myeloid markers, CD206, ASC, and PU.1 (Figure 1 – figure supplement 2). Lyz2Credirected Cre- recombination with 90-100% efficiency in CD206+ cells and with 50-70% efficiency in ASC+ and PU.1+ cells, the range depending on whether tdTomato or GFP colocalization was being scored (Figure 1 – figure supplement 2). In the text that follows, we will refer to Lyz2Cre- recombined cells simply as “myeloid cells”, although they should be understood as CD206+ myeloid cells. The low abundance of TLR4 and the limitations of commercial anti-TLR4 antibodies precluded a direct immunohistochemical assessment of TLR4 loss in Tlr4VEKO and Tlr4MKO mice.

Earlier version: To profile the transcriptional landscape of the leptomeninges in WT, Tlr4VEKO, and Tlr4MKO mice, with and without meningitis, we performed single-nucleus RNA sequencing (snRNA-seq) on dissected leptomeninges (Supplementary file 1). A uniform manifold approximation and projection (UMAP) analysis of the dataset revealed six major cell classes: Fb (fibroblast)-arachnoid1, Fb-pia, arachnoid barrier cells, dural border cells, vascular cells, and immune/myeloid cells (Figure 1C; Figure 1 – figure supplement 3A–C), consistent with prior studies (Pietilä et al., 2023; Wang et al., 2023; Smyth et al., 2024; we note that the cell cluster referred to in Wang et al., 2023 as “fibroblasts dura 3” corresponds to dural border cells). The same set of cell clusters was observed with both PIP-seq (Clark et al., 2023) and 10X Genomics platforms, which additionally confirmed expression of Cdh5 in non-myeloid and non-EC leptomeningeal cell types and expression of Lyz2 exclusively in myeloid cells (Figure 1 – figure supplement 4).

Revised version: To profile the transcriptional landscape of the leptomeninges in WT, Tlr4VEKO, and Tlr4MKOmice, with and without meningitis, we performed single-nucleus RNA sequencing (snRNA-seq) on dissected leptomeninges (Supplementary file 1). A uniform manifold approximation and projection (UMAP) analysis of the dataset revealed six major cell classes: Fb (fibroblast)-arachnoid1, Fb-pia, arachnoid barrier cells, dural border cells, vascular cells, and immune/myeloid cells (Figure 1C; Figure 1 – figure supplement 3A–C), consistent with prior studies (Pietilä et al., 2023; Wang et al., 2023; Smyth et al., 2024; we note that the cell cluster referred to in Wang et al., 2023 as “fibroblasts dura 3” corresponds to dural border cells). The same set of cell clusters was observed with both PIP-seq (Clark et al., 2023) and 10X Genomics platforms, which additionally confirmed expression of Cdh5 in non-myeloid and non-EC leptomeningeal cell types and expression of Lyz2 exclusively in myeloid cells (Figure 1 – figure supplement 4).

Comparison of infected and uninfected conditions showed that Tlr4VEKO mice exhibited a global attenuation of infection-induced transcriptional responses across all major leptomeningeal cell types, as judged by the positions of cell clusters in the UMAP (Figure 1D). Similar trends were observed in principal component analysis (PCA), where infected and uninfected WT samples were widely separated along the PC2 axis, whereas infected and uninfected Tlr4VEKO samples clustered closely together (Figure 1 – figure supplement 5). Uninfected Tlr4MKO transcriptomes localized roughly midway between infected and uninfected WT samples along PC2.

Earlier version: At the level of individual genes, all major leptomeningeal cell types displayed infection-associated transcriptional changes in which the most differentially expressed genes, based on a comparison between infected and uninfected WT mice, exhibited similarly responses in WT and Tlr4MKO mice but little or no response in Tlr4VEKO mice (Figure 1E-G). In WT and Tlr4MKO mice, infection upregulated NF-κB and TNF-α transcriptional responses, whereas Tlr4VEKO mice were largely devoid of these responses (Figure 1H and Figure 1 – figure supplement 3D). Similar trends were observed in dot plots showing transcriptional changes associated with JAK-STAT and Interferon (IFN)-ɣ signaling pathways (Figure 1 – figure supplement 5). Interestingly, leptomeningeal cells in uninfected Tlr4MKO mice exhibit a transcriptional state that partially mimics the infection response (Figure 1D-H and Figure 1 – figure supplements 5 and 6). We speculate that this Tlr4MKO phenotype could represent an increase in the baseline level of innate immune signaling in uninfected Tlr4MKO mice. Comparing snRNA-seq data sets for individual mice and individual genes confirms for ECs and myeloid cells the conclusions obtained from aggregated data (Figure 1 – figure supplement 7). Together, these findings define an E. coli meningitis-associated transcriptional program controlled by TLR4 signaling that drives inflammatory responses in all of the major leptomeningeal cell types (Figure 1I).

Revised version: At the level of individual genes, all major leptomeningeal cell types displayed infection- associated transcriptional changes in which the most differentially expressed genes, based on a comparison between infected and uninfected WT mice, exhibited similarly responses in WT and Tlr4MKOmice but little or no response in Tlr4VEKO mice (Figure 1E-G). In WT and Tlr4MKO mice, infection upregulated NF-κB and TNF-α transcriptional responses, whereas Tlr4VEKOmice were largely devoid of these responses (Figure 1H and Figure 1 – figure supplement 3D). Similar trends were observed in dot plots showing transcriptional changes associated with JAK-STAT and Interferon (IFN)-ɣ signaling pathways (Figure 1 – figure supplement 5). Interestingly, leptomeningeal cells in uninfected Tlr4MKO mice exhibit a transcriptional state that partially mimics the infection response (Figure 1D-H and Figure 1 – figure supplements 5 and 6). We speculate that this Tlr4MKO phenotype could represent an increase in the baseline level of innate immune signaling in uninfected Tlr4MKO mice. Comparing snRNA-seq data sets for individual mice and individual genes confirms for ECs and myeloid cells the conclusions obtained from aggregated data (Figure 1 – figure supplement 7). Together, these findings define an E. coli meningitis-associated transcriptional program controlled by TLR4 signaling that drives inflammatory responses in all of the major leptomeningeal cell types (Figure 1I).

TLR4 signaling drives ICAM1 induction, increased vascular permeability, and myeloid activation during neonatal E. coli infection

To examine the role of TLR4 signaling in meningitis-associated responses within the leptomeninges, we used whole mount imaging to assess markers of vascular permeability and immune activation. Intercellular adhesion molecule-1 (ICAM1) serves as a well-established marker of endothelial activation during infection and inflammation (Bui et al., 2020; Singh et al., 2023). Whole mount leptomeninges from P6 uninfected and E. coli-infected WT, Tlr4VEKO, and Tlr4MKO mice revealed robust ICAM1 upregulation in WT and Tlr4MKO mice, whereas Tlr4VEKO mice showed little change in ICAM1 levels (Figure 2A–B). These protein-level changes mirrored the genotype-specific transcriptional differences detected by snRNA-seq (Figure 2 – figure supplement 1A). TLR4-dependent ICAM1 upregulation was also observed in underlying cortical capillaries (Figure 2 – figure supplement 1B–C). For these and all other immunostaining experiments with mice or with cultured cells, the numbers of mice and the numbers of images used in the analyses are listed in Supplementary file 2.

Non-myeloid TLR4 signaling drives ICAM-1 induction, increased vascular permeability, and myeloid cell activation during neonatal E. coli infection.

(A) Whole mount leptomeninges from P6 WT, Tlr4VEKO, and Tlr4MKO mice, under control or E. coli-infected conditions, immunostained for Cldn5 and ICAM-1 (upper panels) or immunostained for Cldn5 and incubated with fluorescent streptavidin (lower panels) to visualize the fate of intravascular sulfo-NHS-biotin. Scale bars, 20 µm. (B) Quantification of ICAM1 immunostaining relative to WT controls, as shown in (A). (C) Quantification of extravascular streptavidin binding relative to WT controls, as shown in (A). (D) Whole mount leptomeninges from P6 WT and Tlr4VEKO mice, control or E. coli-infected, immunostained for Cldn5 and either ASC or CD206 to visualize myeloid cells. Scale bars, 20 µm. (E) Quantification of ASC+ area (upper) and CD206+ area (lower), as shown in (D). Box plots show the median, interquartile range, and all individual data points. Each data point represents one image from a single mouse, with two locations imaged per mouse. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05).

To assess BBB integrity within the leptomeninges, we performed intraperitoneal (IP) injections of sulfo-NHS-biotin, a low molecular weight intravascular tracer. In this assay, perivascular streptavidin staining serves as a readout of vascular leakage (Wang et al., 2019). WT and Tlr4MKO mice exhibited marked increases in sulfo-NHS-biotin extravasation upon E. coli infection, whereas Tlr4VEKO mice showed no change from the uninfected state (Figure 2A and C). The simplest interpretation of these data is that the site of sulfo-NHS biotin leakage is primarily vascular. However, we cannot exclude some contribution from increased arachnoid barrier permeability.

Earlier version: In addition to these endothelial responses, we examined alterations in myeloid cells within the leptomeninges and the underlying cortex. The ASC protein (Apoptosis-associated speck-like protein containing a CARD; ASC/PYCARD), an inflammasome adaptor that changes its oligomeric state and subcellular localization in response to inflammatory stimuli (Sester et al., 2016; Franklin et al., 2018), localizes to leptomeningeal myeloid cells and brain microglia (Figure 2D). Quantification of the area occupied by ASC+ cells revealed significant increases in infected WT mice within both the leptomeninges and cortex (Figure 2D and E; Figure 2 – figure supplement 1D–E), whereas infected Tlr4VEKO mice exhibited a more modest response. To further characterize myeloid morphology, we stained for CD206 (Figure 2D and E; Figure 2 – figure supplement 1F). Within the leptomeninges, and similar to ASC staining, the area occupied by CD206+ myeloid cells increased in infected WT mice but showed substantially smaller changes in infected Tlr4VEKO mice (Figure 2D–E). In the brain, CD206+ cells were too sparse to reliably quantify (Figure 2 – figure supplement 1F). The increase in area of ASC and CD206 immunostaining in the infected leptomeninges relative to controls largely reflects a shift of myeloid cells from a compact to a more extended morphology (Wang et al, 2023). These findings imply that non-myeloid TLR4 signaling promotes myeloid activation and associated morphologic changes during E. coli meningitis.

Revised version: In addition to these endothelial responses, we examined alterations in myeloid cells within the leptomeninges and the underlying cortex. The ASC protein (Apoptosis-associated speck-like protein containing a CARD; ASC/PYCARD), an inflammasome adaptor that changes its oligomeric state and subcellular localization in response to inflammatory stimuli (Sester et al., 2015; Franklin et al., 2018), localizes to leptomeningeal myeloid cells and brain microglia (Figure 2D). Quantification of the area occupied by ASC+ cells revealed significant increases in infected WT mice within both the leptomeninges and cortex (Figure 2D and E; Figure 2 – figure supplement 1D–E), whereas infected Tlr4VEKO mice exhibited a more modest response. To further characterize myeloid morphology, we stained for CD206 (Figure 2D and E; Figure 2 – figure supplement 1F). Within the leptomeninges, and similar to ASC staining, the area occupied by CD206+ myeloid cells increased in infected WT mice but showed substantially smaller changes in infected Tlr4VEKO mice (Figure 2D–E). In the brain, CD206+ cells were too sparse to reliably quantify (Figure 2 – figure supplement 1F). The increase in area of ASC and CD206 immunostaining in the infected leptomeninges relative to controls largely reflects a shift of myeloid cells from a compact to a more extended morphology (Wang et al, 2023). These findings imply that non-myeloid TLR4 signaling promotes myeloid activation and associated morphologic changes during E. coli meningitis.

Earlier version: Immunostaining for cleaved Caspase-3+ cells in leptomeninges+cortex flat mounts revealed increased numbers of cleaved Caspase-3+ cortical cells within ∼30 um beneath the pial surface following E. coli infection in WT and Tlr4MKO mice, but not in Tlr4VEKO mice (Figure 2 – figure supplement 2A and B). Similar analyses in coronal sections extending ∼750 um into the cortex showed ∼2-fold greater numbers of cleaved Caspase-3+ cortical cells in WT mice compared to Tlr4VEKO mice (Figure 2 – figure supplement 2C and D). In coronal sections, the number of cleaved Caspase-3+ cells in the leptomeninges correlated with the number in the cortex (R=0.79) in both WT mice and Tlr4VEKO mice (Figure 2 – figure supplement 2E). Co-staining with neuronal (NeuN), glial (GFAP), EC (Cldn5), and myeloid (PU.1) markers revealed that cleaved Caspase-3 in the cortex localized to myeloid cells (Figure 2 – figure supplement 2F). Immunostaining for Iba1, a marker of microglial activation, showed increased numbers of cortical Iba1+ cells in both WT and Tlr4VEKO mice in response to infection, with a ∼2-fold greater number in WT mice compared to Tlr4VEKO mice (Figure 2 – figure supplement 3). The infection-dependent appearance of cleaved Caspase-3+ myeloid cells in the cortex could reflect the previously described non-apoptotic role of caspases in immune cell differentiation and activation (Sordet et al., 2002; Santambrogio et al., 2005). Together, these results indicate that non-myeloid TLR4 signaling drives both increased vascular permeability and myeloid cell activation during neonatal E. coli meningitis.

Revised version: Immunostaining for cleaved Caspase-3+ cells in leptomeninges+cortex flat mounts revealed increased numbers of cleaved Caspase-3+ cortical cells within ∼30 um beneath the pial surface following E. coli infection in WT and Tlr4MKO mice, but not in Tlr4VEKOmice (Figure 2 – figure supplement 2A and B). Similar analyses in coronal sections extending ∼750 um into the cortex showed ∼2-fold greater numbers of cleaved Caspase-3+ cortical cells in WT mice compared to Tlr4VEKOmice (Figure 2 – figure supplement 2C and D). In coronal sections, the number of cleaved Caspase-3+ cells in the leptomeninges correlated with the number in the cortex (R=0.79) in both WT mice and Tlr4VEKO mice (Figure 2 – figure supplement 2E). Co- staining with neuronal (NeuN), glial (GFAP), EC (Cldn5), and myeloid (PU.1) markers revealed that cleaved Caspase-3 in the cortex localized to myeloid cells (Figure 2 – figure supplement 2F). Immunostaining for Iba1, a marker of microglial activation, showed increased numbers of cortical Iba1+ cells in both WT and Tlr4VEKO mice in response to infection, with a ∼2-fold greater number in WT mice compared to Tlr4VEKO mice (Figure 2 – figure supplement 3). The infection- dependent appearance of cleaved Caspase-3+ myeloid cells in the cortex could reflect the previously described non-apoptotic role of caspases in immune cell differentiation and activation (Sordet et al., 2002; Santambrogio et al., 2005). Together, these results indicate that non- myeloid TLR4 signaling drives both increased vascular permeability and myeloid cell activation during neonatal E. coli meningitis.

The loss of TLR4 signaling in Tlr4VEKO mice had little or no effect on the clinical course of the E. coli infection. The presence of E. coli in blood cultures and its dissemination to brain, liver, and lung occurred indistinguishably in infected WT and Tlr4VEKO mice (Figure 2 – figure supplement 4A-C). The most obvious clinical manifestation of E. coli infection is a failure to gain weight, an observation consistent across WT, Tlr4VEKO, Tlr4MKO, and Tlr4-/- mice (Figure 2 – figure supplement 4D; Wang et al., 2023).

Claudin-5 redistribution in response to E. coli in leptomeningeal endothelial cells and in bEnd.3 cells

Earlier version: In our previous description of meningitis-associated responses in the leptomeninges, we observed a redistribution of Cldn5 independent of changes in total Cldn5 levels (Wang et al., 2023). Given the central role of Cldn5 in maintaining BBB integrity, we investigated whether TLR4-mediated inflammatory signaling contributed to these changes. Consistent with prior findings (Wang et al., 2023), both WT and Tlr4MKO mice showed an increase in the area occupied by Cldn5 in the leptomeninges following infection, likely referable to both increased vessel diameter and a redistribution of Cldn5 within ECs (Figure 3A-B; Figure 3 – figure supplement 1C). Similar to the effects of global Tlr4 KO (Wang et al., 2023), Tlr4VEKO mice showed minimal redistribution of Cldn5 (Figure 3A-B). The area occupied by Zonula Occludens-1 (ZO-1), a tight junction scaffold protein, showed a modest but statistically significant increase in WT leptomeningeal vessels but no significant change in Tlr4VEKO leptomeningeal vessels during infection (Figure 3 – figure supplement 1A and B). In sum, these data suggests that non-myeloid TLR4 signaling regulates EC tight-junction organization.

Revised version: In our previous description of meningitis-associated responses in the leptomeninges, we observed a redistribution of Cldn5 independent of changes in total Cldn5 levels (Wang et al., 2023). Given the central role of Cldn5 in maintaining BBB integrity, we investigated whether TLR4-mediated inflammatory signaling contributed to these changes. Consistent with prior findings (Wang et al., 2023), both WT and Tlr4MKO mice showed an increase in the area occupied by Cldn5 in the leptomeninges following infection, likely referable to both increased vessel diameter and a redistribution of Cldn5 within ECs (Figure 3A-B; Figure 3 – figure supplement 1C). Similar to the effects of global Tlr4 KO (Wang et al., 2023), Tlr4VEKOmice showed minimal redistribution of Cldn5 (Figure 3A-B). The area occupied by Zonula Occludens-1 (ZO-1), a tight junction scaffold protein, showed a modest but statistically significant increase in WT leptomeningeal vessels but no significant change in Tlr4VEKO leptomeningeal vessels during infection (Figure 3 – figure supplement 1A and B). In sum, these data suggests that non-myeloid TLR4 signaling regulates EC tight-junction organization.

Claudin-5 redistribution in response to E.coli in leptomeningeal endothelial cells and in bEnd.3 cells.

Earlier version: (A) Whole mount leptomeninges from P6 WT, Tlr4VEKO, and Tlr4MKO mice, control or E. coli-infected, immunostained for Cldn5. Scale bars, 20 µm. (B) Quantification of Cldn5+ area in the leptomeninges, as shown in (A). (C) Control bEnd.3 cells and bEnd.3 cells exposed to live E. coli or live group B Streptococcus (GBS) for 1–6 hours were immunostained for Cldn5 and NF-κB. Scale bars, 20 µm. (D) Control bEnd.3 cells and bEnd.3 cells exposed to heat-killed E. coli or heat-killed group B Streptococcus (GBS) for 1–6 hours were immunostained for Cldn5 and NF-κB. Scale bars, 20 µm. (E) Quantification of experiments shown in (C) and (D). First plot, quantification of log10-transformed cytoplasmic-to-plasma membrane (Cyto/PM) intensity ratio for Cldn5 in bEnd.3 cells exposed to live E. coli for the indicated time in hours. Second plot, quantification of log10-transformed nuclear-to-cytoplasmic (Nuc/Cyto) intensity ratio for NF-κB in bEnd.3 cells exposed to live GBS for the indicated time in hours. The third and fourth plots are analogous to the first and second plots, except that bEnd.3 cells were exposed to heat-killed (HK) E. coli or GBS. Box plot in (B) shows median, interquartile range, and all individual data points. Each in vivo data point represents one image from a single mouse, with two locations imaged per mouse. For plots in (E), each data point represents a single bEnd.3 cell analyzed from representative images in three biological replicates. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05).

Revised version: (A) Whole mount leptomeninges from P6 WT, Tlr4VEKO, and Tlr4MKOmice, control or E. coli- infected, immunostained for Cldn5. Scale bars, 20 µm. (B) Quantification of Cldn5+ area in the leptomeninges, as shown in (A). (C) Control bEnd.3 cells and bEnd.3 cells exposed to live E. coli or live group B Streptococcus (GBS) for 1–6 hours were immunostained for Cldn5 and NF-κB. Scale bars, 20 µm. (D) Control bEnd.3 cells and bEnd.3 cells exposed to heat-killed E. coli or heat-killed group B Streptococcus (GBS) for 1–6 hours were immunostained for Cldn5 and NF-κB. Scale bars, 20 µm. (E) Quantification of experiments shown in (C) and (D). First plot, quantification of log10- transformed cytoplasmic-to-plasma membrane (Cyto/PM) intensity ratio for Cldn5 in bEnd.3 cells exposed to live E. coli for the indicated time in hours. Second plot, quantification of log10- transformed nuclear-to-cytoplasmic (Nuc/Cyto) intensity ratio for NF-κB in bEnd.3 cells exposed to live GBS for the indicated time in hours. The third and fourth plots are analogous to the first and second plots, except that bEnd.3 cells were exposed to heat-killed (HK) E. coli or GBS. Box plot in (B) shows median, interquartile range, and all individual data points. Each in vivo data point represents one image from a single mouse, with two locations imaged per mouse. For plots in (E), each data point represents a single bEnd.3 cell analyzed from representative images in three biological replicates. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05).

To directly test the link between inflammatory signaling and Cldn5 redistribution, we turned to cultured mouse brain ECs (bEnd.3), which express both TLR4 and Cldn5 at baseline (Williams et al., 1988; Montesano et al., 1990; Sikorski et al., 1993). Although they are brain- derived, bEnd.3 cells lack many BBB-specific attributes and, therefore, they likely exhibit generalized endothelial responses rather than brain-specific responses to bacterial exposure. Time-course imaging of bEnd.3 cells exposed to living or heat-killed E. coli or Group B Streptococcus (GBS) revealed rapid redistribution of Cldn5 from the plasma membrane to intracellular compartments within one hour of bacterial exposure under all four conditions (Figure 3C, 3D, and 3E first and third plots). Nuclear translocation of NF-κB was greatest following exposure to live E. coli, beginning within 1 hour and progressively increasing over the 6-hour duration of the experiment, consistent with strong TLR4 activation (Figure 3C left half, and 3E second plot). Heat-killed E. coli produced a more modest nuclear translocation of NF- κB, peaking at 3–6 hours (Figure 3D left half, and 3E fourth plot). Despite the robust internalization of Cldn5 in response to either live or heat-killed GBS, the NF-κB response to GBS was minimal, suggesting that TLR2-mediated NF-κB signaling is weaker than TLR4- mediated signaling in bEnd.3 cells (Figure 3C and D right halves, and 3E second and fourth plots). Together, these findings indicate that Cldn5 redistribution is a rapid endothelial response to bacterial exposure that can occur in the absence of detectable NF-κB nuclear translocation, suggesting partial or complete independence from NF-κB–driven transcriptional responses.

TLR4 signaling controls NF-κB activation, tight junction dynamics, and barrier integrity in bEnd.3 cells exposed to E. coli

Added in the revised version: To directly test the role of TLR4 signaling in bEnd.3 cells exposed to E. coli, we generated a CRISPR Tlr4 knockout line (Figure 4). Tlr4 knockout lines were generated using Cas9 and sgTrack-mCherry-mediated sgRNA delivery (see Methods; Hulton et al., 2020). Exons 1 and 3 of the mouse Tlr4 locus were targeted with sgRNAs and clonal lines were isolated by fluorescence-activated cell sorting. Functional loss of TLR4 was confirmed by quantifying NF-κB nuclear translocation following bacterial exposure (Figure 4 – figure supplement 1A). Sanger sequencing confirmed frameshift mutations in all Tlr4 alleles in several exon 3-targeted clones, one of which (Tlr4KO-m1) was selected for further study (Figure 4 – figure supplement 1B and C). Immunostaining for Cldn5 and NF-κB in WT and Tlr4KO bEnd.3 cells showed that deletion of TLR4 prevented both NF-κB nuclear translocation and Cldn5 internalization in response to E. coli (Figure 4A–D).

Removed in the revised version: To directly test the role of TLR4 signaling in bEnd.3 cells exposed to E. coli, we generated a CRISPR Tlr4 knockout line (Figure 4). Tlr4 knockout lines were generated using Cas9 and sgTrack-mCherry-mediated sgRNA delivery (see Methods; Hulton et al., 2020).

TLR4 signaling controls NF-κB activation, tight junction dynamics, and barrier integrity in bEnd.3 cells exposed to E. coli.

Earlier version: (A) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 1 hour or 4 hours, were fixed and immunostained for Cldn5 and NF-κB and stained with DAPI. Scale bars, 20 µm. (B) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 1 hour or 4 hours, were fixed and immunostained for Cldn5 and ZO-1 and stained with CellTrace. Scale bars: 20 µm. (C) Quantification of log10-transformed nuclear-to-cytoplasmic (Nuc/Cyto) NF-κB intensity ratio in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli, as shown in (A). (D) Quantification of log10-transformed cytoplasmic-to-plasma membrane (Cyto/PM) Cldn5 intensity ratio in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli, using ZO-1 localization to define the plasma membrane region, as shown in (B). (E) WT and Tlr4KO bEnd.3 cells were grown to confluence on biotinylated gelatin-coated coverslips, with or without E. coli exposure for 4 hours, were then incubated with fluorescent streptavidin, and were finally fixed and immunostained for ZO-1. Representative images are shown with merged (upper) and streptavidin-only (lower) channels. Scale bars, 20 µm. (F) Quantification of streptavidin+ area in WT and Tlr4KO bEnd.3 cells with or without exposure to E. coli, as shown in (E). Wells coated with gelatin alone or biotinylated gelatin (without cultured cells) served as negative and positive controls, respectively. Box plots show median, interquartile range, and all individual data points. Each data point represents a single cell (C and D) or one image field (F) from representative images in three biological replicates. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05)

Revised version: (A) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 1 hour or 4 hours, were fixed and immunostained for Cldn5 and NF-κB and stained with DAPI. Scale bars, 20 µm. (B) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 1 hour or 4 hours, were fixed and immunostained for Cldn5 and ZO-1 and stained with CellTrace. Scale bars: 20 µm. (C) Quantification of log10-transformed nuclear-to-cytoplasmic (Nuc/Cyto) NF-κB intensity ratio in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli, as shown in (A). (D) Quantification of log10-transformed cytoplasmic-to-plasma membrane (Cyto/PM) Cldn5 intensity ratio in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli, using ZO-1 localization to define the plasma membrane region, as shown in (B). (E) WT and Tlr4KObEnd.3 cells were grown to confluence on biotinylated gelatin-coated coverslips, with or without E. coli exposure for 4 hours, were then incubated with fluorescent streptavidin, and were finally fixed and immunostained for ZO-1. Representative images are shown with merged (upper) and streptavidin-only (lower) channels. Scale bars, 20 µm. (F) Quantification of streptavidin+ area in WT and Tlr4KO bEnd.3 cells with or without exposure to E. coli, as shown in (E). Wells coated with gelatin alone or biotinylated gelatin (without cultured cells) served as negative and positive controls, respectively. Box plots show median, interquartile range, and all individual data points. Each data point represents a single cell (C and D) or one image field (F) from representative images in three biological replicates. Statistical comparisons (p-values) were calculated with the Wilcoxon rank- sum test. ns, not significant (p > 0.05).

Removed in the revised version: Exons 1 and 3 of the mouse Tlr4 locus were targeted with sgRNAs and clonal lines were isolated by fluorescence-activated cell sorting. Functional loss of TLR4 was confirmed by quantifying NF-κB nuclear translocation following bacterial exposure (Figure 4 – figure supplement 1A). Sanger sequencing confirmed frameshift mutations in all Tlr4 alleles in several exon 3-targeted clones, one of which (Tlr4KO-m1) was selected for further study (Figure 4 – figure supplement 1B and C). Immunostaining for Cldn5 and NF-κB in WT and Tlr4KO bEnd.3 cells showed that deletion of TLR4 prevented both NF-κB nuclear translocation and Cldn5 internalization in response to E. coli (Figure 4A–D).

To test whether TLR4 signaling directly regulates endothelial barrier permeability, we used a cell culture assay to quantify barrier permeability across confluent bEnd.3 monolayers (Dubrovskyi et al., 2013). In this assay, bEnd.3 cells were grown to confluence on biotinylated gelatin and breaks in the monolayer of cells were visualized by adding fluorescent streptavidin to the medium. Cldn5-containing tight junctions were observed only in confluent bEnd.3 monolayers (Figure 4 – figure supplement 2A). Control experiments showed that the streptavidin signal depended on coating the plate with biotinylated gelatin and on gaps in the cell monolayer (Figure 4 – figure supplement 2B). When bEnd.3 cells were grown to confluence, the monolayer effectively restricted streptavidin access to the underlying biotinylated gelatin substrate (Figure 4 – figure supplement 2B). Application of this assay to WT and Tlr4KO bEnd.3 monolayers that were exposed to E. coli showed that loss of TLR4 signaling prevented streptavidin access to the underlying biotinylated gelatin substrate and preserved tight-junction integrity (Figure 4E and F). Together, these data indicate that TLR4 signaling in cultured brain ECs drives Cldn5 internalization and increases monolayer permeability in vitro, mirroring TLR4’s role in vivo.

Comparisons of Cldn5 localization with junctional, plasma membrane, and trafficking markers in WT and Tlr4KO bEnd.3 cells with or without E. coli exposure

Earlier version: Next, we examined the subcellular localization and specificity of Cldn5 redistribution in WT and Tlr4KO bEnd.3 cells. We examined several tight junction-associated and integral membrane proteins expressed in bEnd.3 cells – ZO-1, β-catenin, glucose transporter 1 (GLUT1), and platelet endothelial cell adhesion molecule-1 (PECAM1) – to assess changes in their colocalization with Cldn5 during E. coli exposure (Figure 5A–5C). This analysis was conducted using confocal microscopy with a 40X objective; thus, the designation of “colocalization” is at the resolution conferred by this objective. Cellular colocalization was quantified using percent-overlap analysis in ImageJ (see Methods). At baseline, ZO-1 and β-catenin showed an average of ∼75% overlap with Cldn5, consistent with their colocalization at endothelial tight junctions (Figure 5D, Figure 5 – figure supplement 1A). Both ZO-1 and β-catenin remained plasma membrane-localized after 4 hours of E. coli exposure, while Cldn5 was redistributed to cytoplasmic compartments in WT but not Tlr4KO cells (Figure 5A, 5D, Figure 5 – figure supplement 1A).

Revised version: Next, we examined the subcellular localization and specificity of Cldn5 redistribution in WT and Tlr4KObEnd.3 cells. We examined several tight junction-associated and integral membrane proteins expressed in bEnd.3 cells – ZO-1, β-catenin, glucose transporter 1 (GLUT1), and platelet endothelial cell adhesion molecule-1 (PECAM1) – to assess changes in their colocalization with Cldn5 during E. coli exposure (Figure 5A–5C). This analysis was conducted using confocal microscopy with a 40X objective; thus, the designation of “colocalization” is at the resolution conferred by this objective. Cellular colocalization was quantified using percent-overlap analysis in ImageJ (see Methods). At baseline, ZO-1 and β- catenin showed an average of ∼75% overlap with Cldn5, consistent with their colocalization at endothelial tight junctions (Figure 5D, Figure 5 – figure supplement 1A). Both ZO-1 and β- catenin remained plasma membrane-localized after 4 hours of E. coli exposure, while Cldn5 was redistributed to cytoplasmic compartments in WT but not Tlr4KOcells (Figure 5A, 5D, Figure 5 – figure supplement 1A).

Comparisons of Cldn5 localization with junctional, plasma membrane, and trafficking markers in WT and Tlr4KO bEnd.3 cells with or without E. coli exposure.

(A-C) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 4 hours, were fixed and immunostained for Cldn5 and the indicated markers. Scale bars, 20 µm. (D) Quantification of overlap between Cldn5 and β-catenin, GLUT1, PECAM1, and ZO-1, with each data point representing a 100 µm x 100 µm region of interest (ROI). (E) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 4 hours, were fixed and immunostained for Cldn5 and the indicated markers. Scale bars, 20 µm. (F) Quantification of overlap between Cldn5 and EEA1, LAMP2, and Rab7, as in (D). (G) WT and Tlr4KO bEnd.3 cells, either not exposed to E. coli (control) or exposed to E. coli for 4 hours, were fixed and immunostained for Cldn5 and the indicated markers. Scale bars, 20 µm. (H) Quantification of overlap between Cldn5 and Rab11, PDI, and RCAS1, as in (D). Box plots show median, interquartile range, and all individual data points. Each data point represents one image from representative images in three biological replicates. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05).

Earlier version: We next analyzed non-tight junction membrane proteins. The glucose transporter GLUT1 showed relatively little overlap with Cldn5 at baseline and localized primarily to the plasma membrane (Figure 5B). In these images, GLUT1 immunostaining was present across the cell but it did not reveal a silhouette of the nucleus, implying that it was localized to the plasma membrane rather than the cytoplasm. During E. coli exposure, GLUT1 was internalized alongside Cldn5, showing ∼75% overlap within cytoplasmic compartments and revealing the silhouette of the nucleus (Figure 5D, Figure 5 – figure supplement 1A). In contrast, PECAM1, another non-tight junction surface marker, retained its plasma-membrane localization during the 4-hour exposure to E. coli in both WT and Tlr4KO cells (Figure 5C, 5D, Figure 5 – figure supplement 1A). Taken together, these findings reveal distinct specificities of internalization among membrane and membrane-associated proteins dependent on TLR4 signaling in cultured brain endothelial cells.

Revised version: We next analyzed non-tight junction membrane proteins. The glucose transporter GLUT1 showed relatively little overlap with Cldn5 at baseline and localized primarily to the plasma membrane (Figure 5B). In these images, GLUT1 immunostaining was present across the cell but it did not reveal a silhouette of the nucleus, implying that it was localized to the plasma membrane rather than the cytoplasm. During E. coli exposure, GLUT1 was internalized alongside Cldn5, showing ∼75% overlap within cytoplasmic compartments and revealing the silhouette of the nucleus (Figure 5D, Figure 5 – figure supplement 1A). In contrast, PECAM1, another non-tight junction surface marker, retained its plasma-membrane localization during the 4-hour exposure to E. coli in both WT and Tlr4KOcells (Figure 5C, 5D, Figure 5 – figure supplement 1A). Taken together, these findings reveal distinct specificities of internalization among membrane and membrane-associated proteins dependent on TLR4 signaling in cultured brain endothelial cells.

Earlier version: To identify the intracellular compartments involved in Cldn5 trafficking, we co-stained WT and Tlr4KO cells for Early Endosome Antigen 1 (EEA1), Lysosome-associated Membrane Protein 2 (LAMP2), Rab7, and Rab11 to label endosomal and lysosomal vesicles, and for Protein Disulfide Isomerase (PDI) and Receptor-binding Cancer Antigen Expressed on SiSo Cells (RCAS1) to mark the endoplasmic reticulum and Golgi apparatus, respectively (Figure 5E-H) (Stamatovic et al., 2009, 2017; Scott et al., 2014; Shearer and Petersen, 2019; MacDonald et al., 2020). After 4 hours of exposure to E. coli, Cldn5 colocalization with LAMP2⁺ vesicles increased to ∼50% overlap in a TLR4-dependent manner (Figure 5E and 5F). E. coli exposure also produced small increases in Cldn5 colocalization with EEA1 and PDI, to ∼20% overlap, which was similarly dependent on TLR4 (Figure 5E-H). E. coli exposure produced little change in the low-level colocalization of Cldn5 with Rab7, Rab11, or RCAS1, and there was little or no effect of TLR4 loss (Figure 5E-H). We interpret the TLR4-dependent increase in Cldn5-EEA1 and Cdn5-LAMP2 overlap as representing the endocytic internalization of Cldn5 and its delivery to lysosomes. We interpret the modest increase in Cldn5-PDI as likely arising from the close proximity of endosomes and lysosomes to the ER, which is present throughout the cytoplasm.

Revised version: To identify the intracellular compartments involved in Cldn5 trafficking, we co-stained WT and Tlr4KO cells for Early Endosome Antigen 1 (EEA1), Lysosome-associated Membrane Protein 2 (LAMP2), Rab7, and Rab11 to label endosomal and lysosomal vesicles, and for Protein Disulfide Isomerase (PDI) and Receptor-binding Cancer Antigen Expressed on SiSo Cells (RCAS1) to mark the endoplasmic reticulum and Golgi apparatus, respectively (Figure 5E-H) (Stamatovic et al., 2009, 2017; Scott et al., 2014; Shearer and Peterson, 2019; MacDonald et al., 2020). After 4 hours of exposure to E. coli, Cldn5 colocalization with LAMP2⁺ vesicles increased to ∼50% overlap in a TLR4-dependent manner (Figure 5E and 5F). E. coli exposure also produced small increases in Cldn5 colocalization with EEA1 and PDI, to ∼20% overlap, which was similarly dependent on TLR4 (Figure 5E-H). E. coli exposure produced little change in the low-level colocalization of Cldn5 with Rab7, Rab11, or RCAS1, and there was little or no effect of TLR4 loss (Figure 5E-H). We interpret the TLR4-dependent increase in Cldn5-EEA1 and Cdn5-LAMP2 overlap as representing the endocytic internalization of Cldn5 and its delivery to lysosomes. We interpret the modest increase in Cldn5-PDI as likely arising from the close proximity of endosomes and lysosomes to the ER, which is present throughout the cytoplasm.

Further insights into the mechanism of Cldn5 endocytosis came from visualizing the internalization of transferrin, cholera toxin subunit B (CTB), and 10-kDa dextran, well- established tracers for clathrin-mediated endocytosis, caveolin-mediated endocytosis, and macropinocytosis, respectively (Figure 6A; Johnson and Spence, 2010). Prior to E. coli exposure, these tracers and Cldn5 exhibited ∼80% overlap at the plasma membrane (Figure 6B). Following E. coli exposure, reduced overlap with transferrin and CTB, together with persistent overlap with 10-kDa dextran, suggested a macropinocytosis-like mechanism of Cldn5 internalization (Figure 6B). The persistence of intracellular Cldn5 immunostaining with up to at least 6 hours of E. coli exposure, together with washout experiments showing restoration of plasma membrane Cldn5 localization within 3 hours of E. coli removal, suggest little or no degradation of internalized Cldn5 during this time frame (Figure 6C and D). Together, these results indicate that TLR4 signaling triggers a highly specific and reversible internalization of Cldn5, providing a potential mechanistic link between inflammatory activation and increased vascular permeability.

In bEnd.3 cells, pharmacologic inhibition of NF-κB nuclear translocation and signaling does not impede rapid internalization of Cldn5 in response to E. coli exposure.

(A) bEnd.3 cells were pre-incubated for 30 minutes with the indicated fluorescent tracer, then either not exposed to E. coli (control) or exposed to E. coli for 4 hours and then fixed and immunostained for Cldn5. Scale bars, 20 µm. (B) Quantification of % overlap between tracer and Cldn5 from the experiment shown in (B). (C) Cldn5 internalization and recovery during 1 hour of E. coli exposure, followed by washout of E. coli, and then an additional 3 hours of incubation in medium without bacteria. Scale bars, 50 µm. (D) Quantification of the experiment shown in (D). The metric on the y-axis is (cytoplasmic signal – plasma membrane signal)/(cytoplasmic signal + plasma membrane signal). 100% cytoplasmic localization corresponds to 1.0. 100% plasma membrane localization corresponds to -1.0. 50% cytoplasmic localization and 50% plasma membrane localization corresponds to 0.0. (E) The subcellular localization of Lamp2 and Cldn5 were determined in bEnd3 cells in the presence or absence of 1 µM of the IKK inhibitor ACHP before E. coli exposure or after 1 hour of E. coli exposure. Scale bars, 20 µm. (F) Quantification of the experiment shown in (A). Each symbol represents a 100 µm x 100 µm region of interest (ROI). (G) NF-κB translocation from cytoplasm to nucleus was visualized in the presence or absence of 1 µM of ACHP before E. coli exposure or after 1 hour of E. coli exposure. Scale bars, 20 µm. (H) Quantification of the experiment shown in (C). Each symbol represents an individual cell. Statistical comparisons (p-values) were calculated with the Wilcoxon rank-sum test. ns, not significant (p > 0.05).

To test the possible role of NF-κB signaling in Cldn5 internalization, we measured Cldn5 co-localization with Lamp2 in bEnd.3 cells after one hour of exposure to E. coli with or without the IκB kinase inhibitor ACHP (2-Amino-6-[2-(cyclopropylmethoxy)-6-hydroxyphenyl]-4-(4- piperidinyl)-3 pyridinecarbonitrile) at 1 µM (Figure 6E and F). Cldn5 internalization and co- localization with Lamp2 were insensitive to ACHP-mediated inhibition of NF-κB signaling under conditions that inhibited NF-κB nuclear translocation (Figure 6G and H). These data imply that the TLR4-dependent internalization of Cldn5 operates via an arm of the TLR4 pathway that is both rapid and distinct from NF-κB mediated changes in gene expression.

Role of TLR4 signaling in the transcriptional response to E. coli in bEnd.3 cells

Earlier version: To define the TLR4-dependent transcriptional response to E. coli exposure in ECs, we performed bulk RNA-seq on WT and Tlr4KO bEnd.3 cells with or without 3 hours of E. coli exposure (Figure 7). Principal component analysis revealed separation of datasets based on both bacterial exposure and genotype (WT vs. Tlr4KO) (Figure 7 – figure supplement 1A and B). Interestingly, WT and Tlr4KO bEnd.3 cells that were not exposed to E. coli clustered separately, suggesting a basal level of TLR signaling even in the absence of E. coli exposure. Differential expression analysis in WT bEnd.3 cells showed robust induction of proinflammatory transcripts upon E. coli exposure, including Il6, Cxcl2, Tnf, Icam1, and Ccl2 (Figure 7A), whereas this response was largely absent in Tlr4KO bEnd.3 cells (Figure 7B; note the change in axis scales between Figures 7A and 7B). Pathway analysis showed that inflammatory signaling was highly TLR4-dependent: TNF-α signaling via NF-κB, IL6–JAK–STAT signaling, and general inflammatory response pathways showed strong transcriptional activation in WT but not in Tlr4KO bEnd.3 cells (Figure 7C–F).

Revised version: To define the TLR4-dependent transcriptional response to E. coli exposure in ECs, we performed bulk RNA-seq on WT and Tlr4KObEnd.3 cells with or without 3 hours of E. coli exposure (Figure 7). Principal component analysis revealed separation of datasets based on both bacterial exposure and genotype (WT vs. Tlr4KO) (Figure 7 – figure supplement 1A and B). Interestingly, WT and Tlr4KO bEnd.3 cells that were not exposed to E. coli clustered separately, suggesting a basal level of TLR signaling even in the absence of E. coli exposure. Differential expression analysis in WT bEnd.3 cells showed robust induction of proinflammatory transcripts upon E. coli exposure, including Il6, Cxcl2, Tnf, Icam1, and Ccl2 (Figure 7A), whereas this response was largely absent in Tlr4KO bEnd.3 cells (Figure 7B; note the change in axis scales between Figures 7A and 7B). Pathway analysis showed that inflammatory signaling was highly TLR4-dependent: TNF-α signaling via NF-κB, IL6–JAK–STAT signaling, and general inflammatory response pathways showed strong transcriptional activation in WT but not in Tlr4KO bEnd.3 cells (Figure 7C–F).

In bEnd.3 cells, TLR4 signaling plays a central role in the transcriptional response to E. coli exposure.

(A) Volcano plot of differential transcript abundances in WT bEnd.3 cells, with or without exposure to E. coli for 3 hours. The plot shows transcript-level log2 fold change (LFC) on the horizontal axis versus −log10(false discovery rate; FDR) on the vertical axis. Red points mark significantly changed transcripts (FDR < 0.05 and |LFC| ≥ 1); grey points are not significant. (B) Volcano plot of differential transcript abundances in Tlr4KO bEnd.3 cells, with or without exposure to E. coli for 3 hours, plotted as in (A). Note the change of scale of the vertical and horizontal axes between (A) and (B). (C) Pathway effect map from ssGSEA Hallmark scores. Each symbol is a pathway. The horizontal axis is the effect of E. coli exposure for 3 hours on WT bEnd.3 cells (ΔssGSEA = E. coli exposed – not exposed) and the vertical axis is the effect of E. coli exposure for 3 hours on Tlr4KO bEnd.3 cells. Symbol size encodes the within-genotype significance [max(−log10 FDR) across both genotypes], and the color code indicates the genotype x E. coli exposure interaction FDR. The 45-degree diagonal line represents equal effects for the two genotypes. (D–F) ssGSEA pathway activity based on genotype and E. coli exposure for three Hallmark immune response pathways: TNF-α signaling via NF-κB (D), Inflammatory response (E), and IL6–JAK–STAT signaling (F). Symbols represent ssGSEA scores for individual samples. The vertical axis shows the interaction FDR score from a linear model for genotype x E. coli exposure. (G) Transcript abundance changes for 12 immune system genes in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli for 3 hours. (H) Heatmap of transcript abundance changes in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli for 3 hours, for TNF-α signaling via NF-κB. Rows correspond to transcripts with the greatest changes; values are variance-stabilized transformation (VST) z-scores derived from DESeq2-normalized counts. Columns are arranged by genotype and condition. (I) Scatterplot showing transcript abundance changes in WT and Tlr4KO bEnd.3 cells, with or without exposure to E. coli for 3 hours for TNF-α signaling via the NF-κB pathway. Each symbol represents one gene. Red points indicate genes significant for either genotype (FDR < 0.05). The top 15 genes by effect size are labeled. Note the different scales for the horizontal and vertical axes. The diagonal line represents equal effects for the two genotypes.

We next quantified transcripts for representative cytokines, chemokines, and adhesion molecules that define endothelial activation. Il6, Cxcl10, Ccl2, Cxcl5, Cd83, Ccl20, and Icam1 transcripts were all significantly upregulated in WT bEnd.3 cells following E. coli exposure but remained unchanged in Tlr4KO bEnd.3 cells (Figure 7G). Heatmap visualization of the Hallmark “TNF-α signaling via NF-κB” gene set confirmed that NF-κB–responsive genes, including Nfkbia, Ccl2, Il6, and Tnfaip3, were selectively induced in WT cells by E. coli exposure (Figure 7H). Direct fold change comparisons between genotypes showed that induction of NF-κB target genes was strongly attenuated in Tlr4KO cells (Figure 7I).

To assess TLR4-dependent pathway activation following E. coli exposure, we performed single-sample gene set enrichment analysis (ssGSEA) across 50 Hallmark pathways. The pathways most strongly affected by genotype x exposure interactions included IL6–JAK–STAT3 signaling, allograft rejection, inflammatory response, TNF-α signaling via NF-κB, and the interferon-ɣ response (Figure 7 – figure supplement 1C). Heatmaps of Leading-Edge genes in these pathways further highlighted the broad suppression of cytokine and cell adhesion transcript induction in Tlr4KO bEnd.3 cells compared to WT (Figure 7 – figure supplement 1D–F).

Together, the in vivo and bEnd.3 transcriptional analyses identify TLR4 as a dominant regulator of inflammatory gene expression in response to E. coli exposure.

Discussion

This study demonstrates that, in a mouse model of neonatal E. coli meningitis, non- myeloid TLR4 signaling is a key determinant of the inflammatory response in all major leptomeningeal cell types and of the infection-associated increase in vascular permeability. Using conditional Tlr4 knockout models, leptomeningeal single-nucleus RNA sequencing, and Tlr4 knockout in bEnd.3 cells, we show that TLR4 activation directs molecular, cellular, and transcriptional reprogramming of the vascular barrier. Our quantitative analysis of Cdh5- CreER-mediated recombination, showing recombination in all leptomeningeal ECs but in only a minority of other nonmyeloid cells, implies that the blunted infection response in the Tlr4VEKO leptomeninges is likely referable to loss of TLR4 signaling in ECs. In cell culture, endothelial TLR4 activation triggers rapid, selective, and reversible Cldn5 internalization, coincident with increased endothelial permeability. At the transcriptional level, endothelial TLR4 controls NF-κB and IL6/JAK/STAT inflammatory programs and cytokine gene expression. These findings establish TLR4 as the central leptomeningeal sensor of gram-negative bacteria, linking innate immune recognition to barrier dysfunction and to the tissue-wide inflammatory response.

The leptomeningeal response to bacterial infection

Our snRNA-seq analyses revealed a dramatic transcriptome response to neonatal E. coli infection across all major leptomeningeal cell types that was broadly attenuated with non- myeloid deletion of Tlr4. Non-myeloid TLR4 signaling was required for the induction of canonical immune activation markers, including ICAM1 and leukocyte-recruiting chemokines. In contrast, myeloid-specific inactivation of Tlr4 with Lyz2-Cre had minimal effect in the mouse meningitis model. Taken together, these data point to innate immune sensing as the gatekeeper of immune escalation in the leptomeninges. It is possible that myeloid TLR4 signaling in the leptomeninges may augment non-myeloid signaling or become relevant later in the course of the infection. We note that these experiments do not distinguish between local vs. distal anatomic sources of LPS or downstream effector molecules, such as cytokines, as activators of the leptomeningeal inflammatory response (Huang et al., 2021). Histologic observations of RFP-expressing E. coli in the brain, liver, and lungs, together with positive blood cultures, indicate substantial systemic dissemination of E. coli in this model. Thus, the inflammatory responses of leptomeningeal cells likely reflect exposure to bacterial products and inflammatory mediators derived from both local and distal sources.

Earlier version: In bEnd.3 cells, TLR4 activation triggers a rapid (≤1 hour) and reversible internalization of Cldn5, and Tlr4KO in bEnd.3 cells prevents both NF-κB translocation from cytoplasm to nucleus and Cldn5 redistribution. With TLR4 activation, ZO-1 and β-catenin remain junctional, PECAM1 remains at the plasma membrane, and GLUT1 co-internalizes with Cldn5, implying protein-specific re-localization. Percent-overlap and tracer analyses support endocytic routing of Cldn5, with prominent LAMP2+ (i.e., endosomal/lysosomal) localization, in a pattern that argues against canonical clathrin/caveolin pathways and is most consistent with macropinocytosis-like uptake. Rapid Cldn5 internalization occurred independently of NF-κB nuclear translocation, suggesting regulation through a non-transcriptional arm of TLR4 signaling. Strikingly, washout of E. coli restores Cldn5 to cell-cell junctions within three hours. These observations are broadly aligned with prior work showing cytokine/chemokine-evoked endocytosis of tight-junction proteins (including claudins and occludin) during EC barrier opening, via Rho/ROCK or caveolae-dependent routes (Stamatovic et al., 2006, 2009, 2012). Functionally, Cldn5 internalization was accompanied by increased barrier permeability, as measured by streptavidin passage across bEnd.3 monolayers or by Sulfo-NHS-biotin leakage in vivo. These findings reveal an EC-intrinsic process linking TLR4 activation to tight-junction remodeling and increased vascular permeability. In classic BBB tight-junction biology, Cldn5 is considered a gatekeeper of size-selective permeability (Morita et al., 1999). These findings extend that picture by identifying Cldn5 internalization as a disease-induced process that is rapid, selective, and reversible.

Revised version: In bEnd.3 cells, TLR4 activation triggers a rapid (≤1 hour) and reversible internalization of Cldn5, and Tlr4KOin bEnd.3 cells prevents both NF-κB translocation from cytoplasm to nucleus and Cldn5 redistribution. With TLR4 activation, ZO-1 and β-catenin remain junctional, PECAM1 remains at the plasma membrane, and GLUT1 co-internalizes with Cldn5, implying protein-specific re-localization. Percent-overlap and tracer analyses support endocytic routing of Cldn5, with prominent LAMP2+ (i.e., endosomal/lysosomal) localization, in a pattern that argues against canonical clathrin/caveolin pathways and is most consistent with macropinocytosis-like uptake. Rapid Cldn5 internalization occurred independently of NF-κB nuclear translocation, suggesting regulation through a non-transcriptional arm of TLR4 signaling. Strikingly, washout of E. coli restores Cldn5 to cell-cell junctions within three hours. These observations are broadly aligned with prior work showing cytokine/chemokine-evoked endocytosis of tight-junction proteins (including claudins and occludin) during EC barrier opening, via Rho/ROCK or caveolae-dependent routes (Stamatovic et al., 2006, 2009, 2012). Functionally, Cldn5 internalization was accompanied by increased barrier permeability, as measured by streptavidin passage across bEnd.3 monolayers or by Sulfo-NHS-biotin leakage in vivo. These findings reveal an EC-intrinsic process linking TLR4 activation to tight-junction remodeling and increased vascular permeability. In classic BBB tight-junction biology, Cldn5 is considered a gatekeeper of size-selective permeability (Morita et al., 1999). These findings extend that picture by identifying Cldn5 internalization as a disease-induced process that is rapid, selective, and reversible.

Meningitis and vascular permeability

Mouse models of neonatal E. coli meningitis have shown conflicting effects of global TLR signaling on the clinical course of infection. Krishnan et al. (2013) reported that global loss of TLR4 resulted in a more rapid downhill course, whereas Jia et al. (2025) reported that global loss of TLR4 conferred substantial protection. Interestingly, Krishnan et al. (2013) reported that global loss of TLR2 conferred substantial protection in the same early postnatal E. coli meningitis model. The observations reported here are intermediate between those of Krishnan et al. (2013) and Jia et al. (2025), as we observed no obvious differences in two measures of disease progression (bacterial burden and body weight) between infected WT and Tlr4VEKO mice at 24 hours post-infection, despite substantial blunting of infection-associated gene expression and permeability changes in the leptomeninges of Tlr4VEKO mice. These divergent findings may reflect differences in experimental models and/or the fact that global TLR deletion affects multiple compartments simultaneously, including immune cell pathogen sensing, systemic inflammatory responses, and vascular barrier regulation.

Earlier version: Multiple pathways and mechanisms have been proposed to explain the increase in vascular permeability in the leptomeninges following bacterial infection. In an EC monolayer model of E. coli infection, activation of protein kinase C-alpha (PKC-alpha) led to disassembly of VE-cadherin mediated cell-cell junctions and a concomitant increase in monolayer permeability (Sukumaran and Prasadarao, 2003). In a similar EC monolayer model of E. coli infection, NOD-like receptor family pyrin domain containing 6 (NLRP6), a regulator of caspase, NF-κB, and inflammasome signaling, was found to reduce E. coli-induced tight junction disruption while also increasing E. coli-induced cell death (Jia et al., 2026). Wang, et al. (2025a) found that RIPK1, a kinase and scaffolding protein that integrates signals from inflammatory and cell death pathways, plays a central role in endothelial cell death and BBB breakdown during E. coli meningitis. Huang et al. (2024) showed that Wnt/beta-catenin signaling, which controls BBB-specific gene expression and is essential for maintaining BBB integrity, is inhibited by NF-κB signaling; augmenting Wnt/beta-catenin signaling blunted LPS-induced vascular leakage while reducing Wnt/beta-catenin signaling had the opposite effect.

Revised version: Multiple pathways and mechanisms have been proposed to explain the increase in vascular permeability in the leptomeninges following bacterial infection. In an EC monolayer model of E. coli infection, activation of protein kinase C-alpha (PKC-alpha) led to disassembly of VE-cadherin mediated cell-cell junctions and a concomitant increase in monolayer permeability (Sukumaran and Prasadarao, 2003). In a similar EC monolayer model of E. coli infection, NOD-like receptor family pyrin domain containing 6 (NLRP6), a regulator of caspase, NF-κB, and inflammasome signaling, was found to reduce E. coli-induced tight junction disruption while also increasing E. coli-induced cell death (Jia et al., 2026). Wang, et al. (2025a) found that RIPK1, a kinase and scaffolding protein that integrates signals from inflammatory and cell death pathways, plays a central role in endothelial cell death and BBB breakdown during E. coli meningitis. Huang et al. (2024) showed that Wnt/beta-catenin signaling, which controls BBB-specific gene expression and is essential for maintaining BBB integrity, is inhibited by NF-κB signaling; augmenting Wnt/beta-catenin signaling blunted LPS- induced vascular leakage while reducing Wnt/beta-catenin signaling had the opposite effect. Yang et al. (2024) observed IL22-mediated degradation of occludin, ZO-1, and Cldn5 in cultured ECs, and showed that IL22 acts via a VEGF-STAT3 pathway to increase vascular permeability; in a mouse model of E. coli meningitis, this effect was largely abrogated by knocking out Il22. Finally, models of LPS-induced and Klebsiella pneumoniae-induced BBB breakdown implicate a Caspase-11 pathway that leads to the induction of Gasdermin D pores in CNS EC plasma membranes, leading to BBB breakdown and eventual pyroptotic cell death of ECs (Wei et al., 2024). Thus, the emerging picture is one of multiple inflammatory and cell death pathways converging on endothelial hyper-permeability.

Removed in the revised version: Yang et al. (2024) observed IL22-mediated degradation of occludin, ZO-1, and Cldn5 in cultured ECs, and showed that IL22 acts via a VEGF-STAT3 pathway to increase vascular permeability; in a mouse model of E. coli meningitis, this effect was largely abrogated by knocking out Il22. Finally, models of LPS-induced and Klebsiella pneumoniae-induced BBB breakdown implicate a Caspase-11 pathway that leads to the induction of Gasdermin D pores in CNS EC plasma membranes, leading to BBB breakdown and eventual pyroptotic cell death of ECs (Wei et al., 2024). Thus, the emerging picture is one of multiple inflammatory and cell death pathways converging on endothelial hyper-permeability.

We speculate that increased endothelial permeability in the leptomeninges could represent a double-edged response, deleterious when sustained and/or excessive, yet potentially protective when transient, facilitating the regulated passage of immune mediators and immune cells that are needed to combat infection. More generally, increases in vascular permeability may represent a general response to CNS stress or injury. For instance, others have observed reduced BBB integrity in brain regions remote from a focal optic-nerve injury, following transient global ischemia, in ischemic stroke, and in LPS-driven systemic inflammation (Banks et al., 2015; Ju et al., 2018; Smith et al., 2018; Cottarelli et al., 2025).

Activation of TLR4 signaling and the resulting increase in vascular permeability has direct clinical implications for neonatal meningitis. In particular, the data are consistent with a model in which endothelial TLR4 signaling contributes causally to CNS injury by promoting loss of barrier integrity, thereby permitting inflammatory cytokines and bacterial products to diffuse into the CNS parenchyma. These findings align with clinical observations in infants with bacterial meningitis, where elevated CSF cytokines (e.g., IL-6, TNF-α) and albumin level correlate with poor neurologic outcomes (Garges et al., 2006; Xu et al., 2019; Gao and Hu, 2024). By demonstrating the centrality of endothelial-intrinsic TLR4 signaling, our data highlight the endothelium not merely as a passive target but as a primary effector of neuropathology. These observations raise the possibility that selective modulation of endothelial TLR4 or more proximate regulators of Cldn5 trafficking could preserve vascular integrity, offering a potential avenue for neuroprotection in neonatal meningitis.

Limitations of the study

Several limitations of the present study should be noted. While bEnd.3 cells provide a useful endothelial model, they do not exhibit the distinctive structural and molecular complexity of the blood–brain barrier (Garcia-Gallardo and Campbell, 2025). In particular, the bEnd.3 culture system lacks pericytes and astrocytes. Additionally, studying fixed cells and tissues does not permit detailed analyses of molecular and cellular dynamics. Going forward, live cell imaging of CNS ECs in culture (cell lines and/or acutely isolated cells) and in vivo two-photon imaging of the leptomeninges would complement the present studies by revealing dynamic responses in real time. Finally, the present study focused on E. coli K1, the dominant Gram- negative neonatal pathogen. Future work could assess TLR signaling in response to other bacterial pathogens, such as Group B Streptococcus.

Concluding remarks

In sum, this study shows that TLR4 signaling couples innate immune recognition to inflammatory responses in multiple leptomeningeal cell types, inducing endothelial Cldn-5 internalization and increasing vascular permeability. The rapid and reversible nature of the Cldn5 internalization response implies that tight-junction remodeling is a dynamically regulated process. These results highlight potential opportunities to preserve vascular barrier properties by reducing endothelial TLR4 signaling, stabilizing Cldn5 in tight junctions, or dampening IL6– JAK–STAT and NF-κB signaling in endothelial cells (Freitas et al., 2018; Gadina et al., 2018; Zaffaroni and Peri, 2018; Hashimoto et al., 2021).

Added in the revised version: Materials and Methods

Removed in the revised version: Methods

Mouse Models and E. coli Infection

Mice were maintained on a C57BL/6J background. Non-myeloid Tlr4 conditional knockout mice (referred to as Tlr4VEKO) were generated by crossing Tlr4fl/fl (JAX, stock #024872) females with Cdh5-CreER/+;Tlr4−/− males, yielding Cdh5-CreER/+;Tlr4fl/–offspring. Myeloid-specific conditional knockout mice (Tlr4MKO), were generated by crossing Tlr4fl/fl females with Lyz2Cre/+;Tlr4−/−males, yielding Lyz2Cre/+;Tlr4fl/– offspring. We note that for both conditional knockout mouse lines, only a single Cre-mediated recombination event was required to convert Tlr4fl/- heterozygous cells to Tlr4-/- homozygous cells. For both sets of crosses, littermates lacking Cre were used as controls. The Cdh5-CreER line (Monvoisin et al., 2006) was the same line used in Wang. et al. (2025a). The Lyz2Cre line (JAX, stock #004781) was described in Clausen et al. (1999). [For the present study, the Lyz2Cre line was chosen based on the observations of Abram et al (2014) that it gives the most efficient myeloid-specific Cre- recombination among twelve Cre lines tested. We note that McKinsey et al (2020) have found that PF4Cre also recombines LoxP targets efficiently in leptomeningeal myeloid cells.] The Rosa26-LSL-tdT-2A-nlsGFP reporter line (JAX, stock #030867) was described in Wang et al. (2018). The Rosa26-LSL-SUN1-sfGFP reporter line (JAX, stock #021039) was described in Mo et al. (2015). Neonatal mice were genotyped by PCR and used for meningitis experiments without sex selection. Meninges snRNA-seq experiments and histological studies used postnatal day 6 (P6) mice with age-matched controls. To activate Cre-mediated recombination in progeny from the cross with Cdh5-CreER, 4-hydroxytamoxifen (4HT; Sigma-Aldrich H7904) was dissolved at a final concentration of 2 mg/ml in 1:10 vol:vol ethanol:sunflower seed oil (Sigma-Aldrich S5007), stored in aliquots at -80°C, and administered by intraperitoneal (IP) injection of 40-50 µl at P2. Mice were housed and handled according to the approved Institutional Animal Care and Use Committee protocol of the Johns Hopkins Medical Institutions (protocol MO22M375).

Neonatal E. coli K1 infection model

Earlier version: To model neonatal meningitis, postnatal day 5 (P5) mice were used, approximating the developmental stage of the human newborn blood–brain barrier and immune system as described in Wang et al. (2023). Litters were randomized to receive a subcutaneous inoculation in the dorsal skin of 1.2 X 10⁵ colony-forming units (CFU) of E. coli RS218 (O18:K1:H7) with an RFP-expression plasmid (Zhu et al., 2020) in 20 µL PBS or 20 µL PBS alone. Twenty-four hours later (at P6), mice were euthanized for tissue collection. Brains, meninges, and relevant tissues were harvested, fixed, and processed for immunofluorescence and tracer leakage assays.

Revised version: To model neonatal meningitis, postnatal day 5 (P5) mice were used, approximating the developmental stage of the human newborn blood–brain barrier and immune system as described in Wang et al. (2023). Litters were randomized to receive a subcutaneous inoculation in the dorsal skin of 1.2 × 10⁵ colony-forming units (CFU) of E. coli RS218 (O18:K1:H7) with an RFP-expression plasmid (Zhu et al., 2020) in 20 µL PBS or 20 µL PBS alone. Twenty-four hours later (at P6), mice were euthanized for tissue collection. Brains, meninges, and relevant tissues were harvested, fixed, and processed for immunofluorescence and tracer leakage assays.

Antibodies for tissue staining

Earlier version: Leptomeningeal and cortical whole mounts were immunostained with the following primary antibodies: mouse anti-Claudin-5 (Alexa Fluor™ 488, ThermoFisher #352588), mouse anti-ZO-1 (Alexa Fluor™ 594, ThermoFisher #339194), rabbit anti-ICAM-1 (CD54, eBioKAT-1, ThermoFisher #14-0542-85), rat anti-CD206 (Bio-Rad MCA2235), rabbit anti-ASC/TMS1 (Cell Signaling Technology #67824S), rat anti-PECAM-1 (BD Pharmingen #553370), rabbit anti-GLUT1 (ThermoFisher MA5-31960), rabbit anti-cleaved Caspase-3 (Cell Signaling Technology, #9664). Alexa dye-conjugated secondary antibodies were from ThermoFisher.

Revised version: Leptomeningeal and cortical whole mounts were immunostained with the following primary antibodies: mouse anti-Claudin-5 (Alexa Fluor™ 488, ThermoFisher #352588), mouse anti-ZO- 1 (Alexa Fluor™ 594, ThermoFisher #339194), rabbit anti-ICAM-1 (CD54, eBioKAT-1, ThermoFisher #14-0542-85), rat anti-CD206 (Bio-Rad MCA2235), rabbit anti-ASC/TMS1 (Cell Signaling Technology #67824S), rat anti-PECAM-1 (BD Pharmingen #553370), rabbit anti- GLUT1 (ThermoFisher MA5-31960), rabbit anti-cleaved Caspase-3 (Cell Signaling Technology, #9664), rabbit anti-ERG (Cell Signaling Technology, #97249). Alexa dye-conjugated secondary antibodies were from ThermoFisher.

Sulfo-NHS-biotin leakage

For analysis of vascular leakage, mice were injected intraperitoneally with Sulfo-NHS-biotin (30 µl of 20 mg/mL Sulfo-NHS-biotin in PBS per mouse; ThermoFisher 21217) 15 minutes before sacrifice. After IP injection, the tracer rapidly equilibrates into the systemic circulation.

Immunostaining of leptomeninges whole mounts

For whole mount imaging of the leptomeninges and adjacent cortex, a preparation was developed to preserve these structures in their native configuration. After euthanasia, brains (with leptomeninges attached) were fixed by immersion in 1% paraformaldehyde in PBS overnight at 4°C. The forebrain was then dissected and hemisected, and a thin slice of tissue from the surface of the brain was obtained by making a 300 µm vibratome section parallel to the surface of the cortex. Tissue slices were permeabilized in 1% Triton X-100/PBS before immunostaining, and blocked with 5% goat or donkey serum, depending on the secondary antibody host species. Primary antibodies were applied at 1:400 dilution, followed by fluorophore-conjugated secondary antibodies. Vascular leakage in mice injected with Sulfo- NHS-biotin was visualized by staining with streptavidin–Alexa Fluor 488 (Thermo Fisher, #S11223).

Confocal microscopy and image acquisition

All images were acquired on a Zeiss LSM 780 laser-scanning confocal microscope controlled by Zeiss Zen software using 20x, 40x, or 63x/1.4 NA oil-immersion objectives. Fluorophores were excited with 405 nm, 488 nm, 561 nm, and 633 nm laser lines, and channels were collected sequentially to minimize spectral overlap. Z-stacks were acquired with step sizes optimized for each staining set; thinner stacks were used for cultured cell monolayers and deeper stacks for leptomeningeal whole mounts. The pinhole was set to 1 Airy unit. Laser power, gain, and offset settings were maintained at constant values within each experimental cohort (e.g. comparisons across genotypes). For leptomeninges imaging, two fields were imaged per hemisphere. Maximum intensity projections or single optical sections were exported from Zen for quantification and figure preparation.

Blood bacterial burden (CFU assay)

Neonatal mice (P6; 24 hours post-infection) were euthanized on ice and whole blood was collected by cardiac puncture into tubes containing EDTA to prevent coagulation. Samples were kept on ice, and 20 µL of whole blood was diluted in 100 µL of LB broth, briefly vortexed, and plated onto LB agar plates. Plates were incubated overnight at 37°C, and colonies (which were visibly red due to RFP expression) were counted the following day. Bacterial burden was reported as colony-forming units (CFUs) per 20 µL of blood. Plates with colony numbers exceeding the reliable counting range were assigned a value of 2,000 CFUs.

Quantifying RFP-expressing E. coli in brain sections

Sections were stained with DAPI to label nuclei and imaged for RFP-expressing E. coli. Fluorescence images were acquired as maximum intensity projections under identical imaging settings across experimental groups. For brain analyses, four anatomically matched territories from coronal sections of cortex were imaged and quantified per mouse. Quantification was performed using fixed threshold-based image analysis in Fiji, where the RFP signal was segmented to generate binary masks representing bacterial area. To define the corresponding cortical tissue area, DAPI images were thresholded and morphologically processed (including smoothing and hole filling) to generate continuous masks of the imaged cortical region. RFP signal was normalized to this DAPI-defined cortical area and reported as the ratio of RFP- positive area to DAPI-defined cortical area.

bEnd.3 cell culture and maintenance

The mouse brain endothelial cell line bEnd.3 (ATCC CRL-2299) was cultured according to ATCC recommendations. Cells were maintained in Dulbecco’s Modified Eagle Medium (DMEM; high glucose, 4.5 g/L) supplemented with 10% fetal bovine serum (FBS). Cultures were kept at 37°C in a humidified 5% CO₂ incubator. Cells were passaged at ∼80–90% confluence using 0.05% trypsin–EDTA and reseeded at a split ratio of 1:4–1:6.

CRISPR-Cas9 generation of bEnd.3 Tlr4KO cell lines

A bEnd.3-Cas9 monoclonal line was first established by transducing bEnd.3 cells with a lentivirus carrying Streptococcus pyogenes Cas9 (Addgene #52962) (Sanjana et al., 2014). Transduced cells were serially diluted to generate monoclonal populations, and Cas9 expression was validated by western blot. Using this Cas9-expressing line, Tlr4 knockouts were generated by transduction with sgRNAs targeting exon 3 of Tlr4 delivered with the sgTrack-mCherry lentivirus (Addgene #114013) (Hulton et al., 2020). Transduced cells were sorted by flow cytometry for mCherry expression, seeded in 96-well trays at one cell per well, and individual colonies were expanded. Genomic DNA was analyzed by PCR amplification across the mutated region in Tlr4 exon 3, followed by subcloning of PCR products into a plasmid vector and sequencing of individual plasmids following colony PCR. A cloned bEnd.3 derivative cell line was considered a candidate Tlr4KO if none of its cloned PCR products showed the WT sequence and all cloned PCR products showed frameshift mutations in Tlr4 exon 3. Functional loss of TLR4 was confirmed by quantifying NF-κB nuclear translocation after E. coli exposure compared to parental controls.

bEnd.3 immunofluorescence

For all experiments involving responses to bacterial exposure, bEnd.3 cells were seeded onto glass coverslips at ∼80% confluency two days prior to the start of the experiment to allow them to form a confluent monolayer with tight junctions. After exposure to E. coli or GBS or mock exposure, cells were fixed in 1% paraformaldehyde in PBS for 15 min at room temperature, permeabilized in 0.1% Triton X-100/PBS, and blocked with 5% goat or donkey serum depending on the secondary antibody host species. Primary antibodies were applied at 1:400 dilution overnight at 4°C, followed by fluorophore-conjugated secondary antibodies. Coverslips were mounted in antifade medium and imaged by confocal microscopy.

Antibodies for bEnd.3 cell immunostaining

Cultured bEnd.3 cells were immunostained with the following primary antibodies: mouse anti- Claudin-5 (Alexa Fluor™ 488, Thermo Fisher #352588), mouse anti-ZO-1 (Alexa Fluor™ 594, Thermo Fisher #339194), rabbit anti-NF-κB p65 (Cell Signaling Technology #8242S), rabbit anti-β-catenin (Cell Signaling Technology #8480), rat anti-PECAM-1 (Thermo Fisher #63-0311- 82), rabbit anti-GLUT1 (Thermo Fisher MA5-31960), mouse anti-EEA1 (Cell Signaling Technology #48453T), rabbit anti-LAMP2 (Cell Signaling Technology #34141T), rabbit anti- Rab7 (Cell Signaling Technology #9367T), rabbit anti-Rab11 (Cell Signaling Technology #5589T), rabbit anti-PDI (Cell Signaling Technology #3501T), and rabbit anti-RCAS1 (Cell Signaling Technology #12290T). Primary antibodies were diluted 1:400 in PBS with 1% Triton X-100 and applied overnight at 4 °C, followed by fluorophore-conjugated secondary antibodies (Thermo Fisher). Coverslips were mounted and imaged by confocal microscopy.

bEnd.3 streptavidin leak assay

Earlier version: Barrier permeability was assayed using biotin-conjugated gelatin substrates as described by Dubrovskyi et al. (2013). Bovine gelatin (Sigma, G1393) was diluted to 1% in D-PBS, prewarmed to 37°C, and dissolved with stirring in a 70°C water bath. EZ-Link NHS-LC-Biotin (Thermo Fisher, Cat. No. 21336) dissolved in DMSO (5 mg/mL) was added to the gelatin at a final concentration of 0.5 mg/mL and reacted for 1 hour at room temperature with constant stirring. The solution was clarified by centrifugation (10,000 x g, 5 minutes), aliquoted, and stored at –20°C. For coating, biotinylated gelatin was thawed, diluted in 0.1 M bicarbonate buffer (pH 8.3) to 0.25 mg/mL, sterilized by 0.22 µm filtration, and applied to culture substrates (e.g., 50 µl per well of a 96-well plate, 2 ml per 35 mm dish). Adsorption was performed overnight at 4°C, followed by PBS washing. Unconjugated gelatin was used as a control. Incorporation of biotin into gelatin was validated using the Pierce Biotin Quantitation Kit (Thermo Fisher, Cat. No. 28005) according to the manufacturer’s instructions. bEnd.3 monolayers were grown to confluence on coated coverslips, and control experiments included (i) gelatin without cells, (ii) biotinylated gelatin without cells, (iii) low-density cells on gelatin, and (iv) low- or high-density cells on biotinylated gelatin. To quantify the integrity of the bEnd.3 monolayer, streptavidin–FITC tracer (Thermo Fisher, #11-4317-87) at 2 µg/mL in DMEM was added to the cell culture medium for 5 min, followed by extensive PBS washing and fixation in 1% paraformaldehyde. Samples were immunostained for ZO-1, nuclei were visualized with DAPI, and the coverslips were mounted in antifade solution and imaged by confocal microscopy. Quantification of streptavidin signal was measured as area fraction per field of view. Each condition was repeated in three independent biological replicates, with multiple image fields analyzed per replicate. To ensure comparability across conditions, all images were collected using identical confocal acquisition settings, and exposure times and intensity scaling were held constant within each experiment.

Revised version: Barrier permeability was assayed using biotin-conjugated gelatin substrates as described by Dubrovskyi et al. (2013). Bovine gelatin (Sigma, G1393) was diluted to 1% in D-PBS, prewarmed to 37°C, and dissolved with stirring in a 70°C water bath. EZ-Link NHS-LC-Biotin (Thermo Fisher, Cat. No. 21336) dissolved in DMSO (5 mg/mL) was added to the gelatin at a final concentration of 0.5 mg/mL and reacted for 1 hour at room temperature with constant stirring. The solution was clarified by centrifugation (10,000 x g, 5 minutes), aliquoted, and stored at –20°C. For coating, biotinylated gelatin was thawed, diluted in 0.1 M bicarbonate buffer (pH 8.3) to 0.25 mg/mL, sterilized by 0.22 µm filtration, and applied to culture substrates (e.g., 50 µl per well of a 96-well plate, 2 ml per 35 mm dish). Adsorption was performed overnight at 4°C, followed by PBS washing. Unconjugated gelatin was used as a control.

Added in the revised version: Incorporation of biotin into gelatin was validated using the Pierce Biotin Quantitation Kit (Thermo Fisher, Cat. No. 28005) according to the manufacturer’s instructions. bEnd.3 monolayers were grown to confluence on coated coverslips, and control experiments included (i) gelatin without cells, (ii) biotinylated gelatin without cells, (iii) low-density cells on gelatin, and (iv) low- or high-density cells on biotinylated gelatin. To quantify the integrity of the bEnd.3 monolayer, streptavidin–FITC tracer (Thermo Fisher, #11-4317-87) at 2 µg/mL in DMEM was added to the cell culture medium for 5 min, followed by extensive PBS washing and fixation in 1% paraformaldehyde. Samples were immunostained for ZO-1, nuclei were visualized with DAPI, and the coverslips were mounted in antifade solution and imaged by confocal microscopy. Quantification of streptavidin signal was measured as area fraction per field of view. Each condition was repeated in three independent biological replicates, with multiple image fields analyzed per replicate. To ensure comparability across conditions, all images were collected using identical confocal acquisition settings, and exposure times and intensity scaling were held constant within each experiment.

bEnd.3 and E. coli washout experiments

bEnd.3 cells were exposed to E. coli K1 (at a 5:1 ratio of E.coli:bEnd.3 cells) in serum- containing DMEM for 1 hour at 37°C, 5% CO₂. For washout, the medium was removed, and cells were washed extensively with pre-warmed PBS, then returned to fresh complete medium for a 3-hour recovery period. Parallel groups were fixed at 1 hour (no washout) or after the 3- hour washout. Cells were fixed in 1% paraformaldehyde, permeabilized, blocked, and immunostained (e.g., for Cldn5 and ZO-1) as described above, then mounted for confocal imaging. Acquisition settings (laser power, gain, offset, and channel order) were held constant across conditions within each experiment. Quantification of Claudin-5 cytoplasm/plasma membrane ratios followed the CellProfiler pipeline (Stirling et al., 2021), and results were analyzed as log10-transformed single-cell values with statistical significance assessed using the Wilcoxon rank-sum test. Three independent biological replicates were performed, with multiple fields analyzed per replicate. For inhibition of NF-κB signaling, ACHP (2-Amino-6-[2- (cyclopropylmethoxy)-6-hydroxyphenyl]-4-(4-piperidinyl)-3 pyridinecarbonitrile; Tocris #4547) was added to the cell culture medium at a final concentration of 1 µM.

bEnd.3 tracer experiments

This procedure was performed as described (Stamatovic et al., 2006, 2017). Briefly, bEnd.3 monolayers on glass coverslips were preloaded with plasma-membrane–binding/uptake tracers prior to bacterial exposure: Cholera Toxin Subunit B, Alexa Fluor™ 647 conjugate (CTB-AF647; Thermo Fisher, C34778), Transferrin from human serum, Alexa Fluor™ 647 conjugate (Tf- AF647; Thermo Fisher, T23366), or Fixable Dextran, 10,000 MW, Alexa Fluor™ 647 (Dextran- AF647; Thermo Fisher, D22914). Tracers were added in ice-cold complete medium for 30 minutes at 4°C to allow accumulation at the plasma membrane. Excess tracer was removed by gentle washes with PBS, and cells were returned to tracer-free complete medium at 37°C. Cultures were then exposed to E. coli K1 (at a 5:1 ratio of E.coli:bEnd.3 cells) for the indicated intervals (e.g., 4 hours) or left unexposed as controls. Cells were fixed in 1% paraformaldehyde, permeabilized, blocked, and immunostained (e.g., for Cldn5 and ZO-1) as described above. Colocalization/overlap with Cldn5 was quantified in Fiji using fixed 100 x 100 µm ROIs and AND-masking of binarized channels (see “Percent overlap quantification”), with all acquisition parameters (laser power, gain, offset, channel order) held constant within each experiment. Three independent biological replicates were performed, with multiple fields analyzed per replicate.

E. coli and GBS preparation and exposure of bEnd.3 cells

Escherichia coli K1 (RS218) and its RFP-expressing derivative (ampicillin-resistant) were cultured in LB broth and on LB agar plates with or without ampicillin, depending on the strain. Streptococcus agalactiae (GBS, ATCC 13813) was cultured in brain heart infusion (BHI) broth and on BHI agar plates. For tissue culture experiments, a single colony was grown overnight in the appropriate broth, and a 200 µl aliquot of the overnight culture was diluted into fresh medium and grown for 1 hour at 37°C. Bacterial density was estimated by optical density at 600 nm, and bacteria were added to cultured bEnd.3 cells (at a 5:1 ratio of bacteria to bEnd.3 cells for all experiments). Heat-killed bacteria were prepared in parallel by heating suspensions to 95°C for 30 min, and bacterial killing was confirmed by plating aliquots to check for the absence of growth. For bacterial exposure experiments, bEnd.3 cells were cultured in the absence of antibiotics in serum-containing medium, and bacteria remained in the medium for the duration of the experiment, except in designated washout assays.

Image Analysis: Iba1 quantification

Immunostained Iba1 image quantification was performed as described above for quantifying RFP-expressing E. coli in brain sections, except that for Iba1 the superficial cortical territory beneath the leptomeninges was manually delineated in coronal sections based on the pattern of Cldn5 immunostaining. Iba1-positive signal was quantified using fixed threshold-based segmentation and normalized to the defined cortical area.

Image Analysis: nuclear-to-cytoplasmic (Nuc/Cyto) intensity ratio

NF-κB and Ki-67 nuclear translocation was quantified using a custom CellProfiler pipeline (Stirling et al., 2021). Nuclei were segmented by DAPI, cytoplasmic regions were defined as the cell body excluding the nucleus, and mean intensities were measured in both compartments. Ratios of nuclear-to-cytoplasmic (Nuc/Cyto) intensity were computed per cell and log10-transformed for analysis. Output spreadsheets from CellProfiler were imported into R (v4.3) and annotated by parsing file names to assign genotype, clone, treatment, and timepoint metadata. Data from replicate experiments were combined, and box plots were generated to show distributions of single-cell values. Statistical comparisons between conditions were performed using the Wilcoxon rank-sum test. For validation, replicate-level plots were generated, outliers were defined by interquartile range (1.5 x IQR) and excluded from simplified representative plots, and a subset of clones and timepoints (e.g., 0, 1, and 3 hours) were displayed for final figures.

Image Analysis: cytoplasmic-to-plasma membrane (Cyto/PM) intensity ratio

Claudin-5 redistribution between the plasma membrane and cytoplasm was quantified in CellProfiler (Stirling et al., 2021) using a custom pipeline. Plasma membrane regions were segmented using ZO-1 as a boundary, while cytoplasmic regions were defined as the cell body excluding the nucleus and plasma membrane masks. Integrated mean intensities of Claudin-5 were measured in both regions, and cytoplasmic-to-plasma membrane (Cyto/PM) ratios were calculated for each cell and log₁₀-transformed for analysis. Output spreadsheets from CellProfiler were imported into R (v4.3) for downstream analysis. File names were parsed to assign genotype, treatment, timepoint, and replicate metadata, and data from replicate experiments were combined. Box plots were generated to display distributions of single-cell values, with statistical comparisons performed using the Wilcoxon rank-sum test. For representative figures, a single replicate per group was shown, and outliers were removed based on the interquartile range (1.5 x IQR).

Image Analysis: percent overlap quantification (ImageJ)

Colocalization between Cldn5 and trafficking/endothelial markers (e.g., EEA1, Rab7, Rab11, LAMP2, ZO-1) or tracers was quantified in Fiji (Schindelin et al., 2012) using a scripted workflow. For each field, maximum-intensity–projected channels were processed as follows: images were Gaussian-blurred (σ = 2), converted to 8-bit, and binarized using fixed thresholds chosen per channel and experiment. Binary masks were generated with “BlackBackground” enabled and saved for Cldn5 and the comparison marker. Overlap masks were created with Image Calculator (logical AND) between the Cldn5 mask and the marker mask (and, when needed, between Cldn5 and ZO-1). An ROI set comprising nine 100 x 100 µm squares was loaded via the Region of Interest (ROI) Manager and applied uniformly to all images. For each ROI, we recorded raw mean intensities (non-binary images) and, on the binary masks, measured area and mean intensity. “Signal area” was computed as (ROI area x mask mean)/255 for Cldn5 and for the overlap image, providing an area-weighted signal metric. Percent overlap with Cldn5 was then calculated per ROI as 100 x (overlap signal area ÷ Cldn5 signal area). Identical steps were repeated for each endosomal marker or tracer by substituting the corresponding channel and threshold. Outputs were written to CSV and pooled across images/replicates for plotting; statistical comparisons were performed using Wilcoxon rank-sum tests as described in the Statistics section. All acquisition and analysis parameters (thresholds, ROIs, and export settings) were held constant within a given experiment.

Image Analysis: statistics

All image quantifications were performed on raw data exported directly from Fiji or CellProfiler without manual exclusion of individual cells or ROIs, unless otherwise noted for representative plots. For Cyto/PM and Nuc/Cyto ratios, values were log₁₀-transformed prior to statistical testing. Percent overlap measurements were expressed as the fraction of Cldn5-positive signal area colocalizing with the comparison marker within fixed 100 x 100 µm ROIs. For extravascular tracer (Sulfo-NHS-biotin) and ICAM-1 analyses, fluorescence intensity or area fractions were normalized to the ROI area or expressed as fold change relative to WT uninfected controls. For all datasets, each dot represents a single measurement (one cell in vitro or one image field in vivo), and biological replicates correspond to independent mice or independent cell culture experiments. Statistical comparisons between two groups were performed using the two-sided Wilcoxon rank-sum test (Mann–Whitney U). Where multiple groups were compared, pairwise Wilcoxon rank-sum tests were performed, and p-values are reported in figure panels. Box plots display the median and interquartile range, with all individual data points overlaid. All statistical analyses were conducted in R (v4.3) using ggplot2 and ggpubr packages, with significance set at p < 0.05.

snRNA-seq libraries and sequencing

For single-nucleus RNA sequencing, leptomeninges were collected as previously described (Wang et al., 2023). In brief, P6 mice were euthanized 24 hours after infection, brains were rapidly removed, and the leptomeninges were dissected into ice-cold Hibernate-A medium. The leptomeninges were frozen at -80°C.

For the PIP-seq protocol, single nucleus (sn)RNA-seq libraries were prepared using the PIP- seq T20 3’ Single Cell RNA Kit v4.0 PLUS (Fluent BioSciences; Clark et al., 2023). Frozen tissue was suspended in homogenization buffer (0.25 M sucrose, 25 mM KCl, 5 mM MgCl2, 20 mM Tricine-KOH,pH 7.8) supplemented with 1 mM DTT, 0.15 mM spermine, 0.5 mM spermidine, EDTA free protease inhibitor (Roche 11836 170 001), 0.5% IGEPAL-630, and 40 U/mL Protector RNase Inhibitor (Sigma, 03335402001). Samples were homogenized in a 2 mL Dounce homogenizer, using 15 strokes with a loose-fitting pestle followed by 30 strokes with a tight-fitting pestle. The sample was filtered through a 10 μm filter (CellTrix, Sysmex, 04-004- 2324), and nuclei were pelleted in Low Retention 1.5 mL microcentrifuge tubes (Thermo Scientific, 3451) for 5 min at 500 x g at 4 °C. Nuclei were washed twice with Nuclei Suspension Buffer (supplied in the Fluent BioSciences kit) supplemented with 1% BSA and 40 U/mL Protector RNase Inhibitor. Nuclei were counted, and ∼40,000 nuclei (for the T20 kit) or ∼20,000 nuclei (for the T10 kit) were used for library production following the Fluent BioSiences protocol. The resulting snRNA-seq libraries were sequenced on an Illumina NovaSeq X Plus sequencer (Supplementary file 1).

For the 10X genomics protocol, snRNA-seq libraries were constructed from suspended nuclei (prepared as described above) using the 10x Genomics Chromium single-cell 3’ v3 kit following the manufacturer’s protocol. Libraries were sequenced on an Illumina NovaSeq 6000.

Analysis of snRNA-seq data

Reads were aligned with the PIPseeker program (Fluent BioSiences, version 3.3.0) using a STAR index based on the GRCm38 mouse genome with addition of Cre sequences (STAR version 2.7.9a). The CellBender program was used to detect empty droplets and remove background (version 0.3.2; Fleming et al., 2023). The SOLO program was used to remove doublets (scVI-tools 1.2.2-post2) (Bernstein et al., 2020). Manual curation was used to further remove potentially mixed (doublet) nuclei. Clusters representing various brain cells were removed from the final data set. The data were normalized using a regularized negative binomial regression algorithm implemented with the SCTransform function (Hafemeister and Satija, 2019). UMAP dimensional reduction was performed using the R uwot package (https://github.com/jlmelville/uwot) integrated into the Seurat R package (Melville, 2022). Data for the various scatter plots were extracted using the Seurat AverageExpression function, and differential gene expression was analyzed using the Seurat FindMarkers function. The Wilcoxon rank sum test was used to calculate p-values. The p-values were adjusted with a Bonferroni correction using all genes in the dataset. Data exploration, analysis, and plotting were performed using RStudio (RStudio Team, 2020), the tidyverse collection of R packages (Wickham, 2017), and ggplot2 (Wickham, 2009).

Dotplots were generated with the default settings, including default normalization. For Figure 1E-H and Figure 1 – figure supplement 5, differential expression between infected and control samples was calculated using Seurat FindMarkers function using only the WT data. The results were filtered for statistically significant difference per cell type. Differentially expressed genes from the indicated Hallmark gene sets are presented in the dot plots for individual cell clusters.

GSEA analysis

For Gene Set Enrichment Analysis (GSEA; Subramanian et al., 2005), genes were ranked by the fold expression change between control and mutant datasets. The ranked gene list was used to detect enriched gene sets within the Broad Institute Hallmark Gene Sets using the fgsea R package (https://github.com/ctlab/fgsea; Korotkevich et al., 2019).

RNA-seq on WT and Tlr4KO bEnd.3 cells: libraries and sequencing

bEnd.3 mouse brain endothelial cells were cultured in 10-cm tissue culture dishes and grown to confluency under standard conditions (37°C, 5% CO₂) and either exposed or not exposed to E. coli for 3 hours. Total RNA was isolated using the Zymo Research RNA purification kit according to the manufacturer’s instructions. Briefly, culture medium was aspirated, cells were rinsed once with ice-cold PBS, and 600 µL of the provided lysis buffer with β-mercapto-ethanol was added directly to the dish. The lysate was collected by scraping and transferred to a microcentrifuge tube, briefly mixed, and applied to the Zymo spin column. After binding, the column was washed with the supplied wash buffers, including an on-column DNase I treatment step to remove contaminating genomic DNA, and RNA was eluted in 50 µL of DNase/RNase- free water. RNA concentration and purity were assessed by spectrophotometry (A₂₆₀/A₂₈₀), and samples were stored at −80°C. Library construction was performed with the Illumina mRNA stranded kit following the manufacturer’s protocol, and sequencing was performed on an Illumina NovaSeqX with a read length of 300 nucleotides.

RNA-seq on WT and Tlr4KO bEnd.3 cells: RNA-seq processing and normalization

Gene-level count matrices from bEnd.3 cells were analyzed in R (v4.5) with Bioconductor (v3.21). The study included two genotypes (WT and Tlr4KO) under two conditions (control and exposed to E. coli for 3 hours), with biological replication per group (WT: 3 control samples and 4 exposed to E. coli samples; Tlr4KO: 3 control samples and 4 exposed to E. coli samples). Counts were normalized using DESeq2’s median-of-ratios size-factor method to control for library size/composition effects. A variance-stabilizing transform (VST) was calculated for quality control and visualization only (e.g., PCA, sample-distance heatmaps). All samples were retained for analysis after QC review.

RNA-seq on WT and Tlr4KO bEnd.3 cells: Differential expression

To quantify infection responses within each genotype, DESeq2’s negative-binomial generalized linear model was fit separately to WT and in Tlr4KO data and then control vs. exposed to E. coli data were compared. The reported log2 fold changes are the control vs. exposed to E. coli estimates from these genotype-specific analyses. P-values were adjusted for multiple testing with the Benjamini–Hochberg procedure. False-discovery rates (FDR or q-values) are cited throughout. Gene-level WT vs. Tlr4KO effect comparisons were visualized by plotting the two genotype-specific fold changes against one another.

RNA-seq on WT and Tlr4KO bEnd.3 cells: Pathway activity scoring and statistics

Pathway analyses used the MSigDB Hallmark collection for mouse, retrieved via msigdbr. Single-sample gene-set enrichment (ssGSEA) was performed with GSVA to obtain an activity score for each pathway in each sample. Scores were computed from log2-scaled normalized counts (log2 (normalized counts + 1)), using only genes present in both the RNA-seq data and the gene set. For statistical inference, we modeled pathway scores with a linear model that included genotype, infection status, and their interaction. From this model we reported three quantities per pathway: the E. coli exposure effect in WT, the E. coli exposure effect in Tlr4KO, and the genotype x exposure interaction (difference-in-response between genotypes). P-values were adjusted across pathways with the Benjamini–Hochberg procedure. Adjusted values are reported as q-values (FDR).

RNA-seq on WT and Tlr4KO bEnd.3 cells: Visualization

Earlier version: Heatmaps were drawn from VST values after centering each gene by its mean (raw z-scores). Within each Hallmark pathway, genes were ranked by the magnitude of the E. coli exposure effect (difference of condition means within each genotype), and the top genes were displayed to show the pattern and directionality across groups. Boxplots summarize ssGSEA scores by group. Per-facet annotations show FDRs for WT and Tlr4KO infection effects and the interaction term. Figure 7C plots the pathway-level infection responses in WT (x-axis) versus Tlr4KO(y-axis), with point size reflecting within-genotype significance and point color indicating interaction FDR. All figures were exported as vector PDFs.

Revised version: Heatmaps were drawn from VST values after centering each gene by its mean (raw z-scores). Within each Hallmark pathway, genes were ranked by the magnitude of the E. coli exposure effect (difference of condition means within each genotype), and the top genes were displayed to show the pattern and directionality across groups. Boxplots summarize ssGSEA scores by group. Per-facet annotations show FDRs for WT and Tlr4KOinfection effects and the interaction term. Figure 7C plots the pathway-level infection responses in WT (x-axis) versus Tlr4KO(y- axis), with point size reflecting within-genotype significance and point color indicating interaction FDR. All figures were exported as vector PDFs.

RNA-seq on WT and Tlr4KO bEnd.3 cells: Software

Analyses used DESeq2, limma (for pathway-score modeling), SummarizedExperiment, GSVA, msigdbr, pheatmap, ggplot2, and ggrepel within the R/Bioconductor ecosystem.